# Homework 01

In this exercise, you will warm up / refresh your Python skills by analyzing some real data inside of a Jupyter notebook environment. You should write whatever **Python** code you like to solve these problems. You can search on the web for documentation and techniques to help you solve the problem. See especially the official [Python tutorial](https://docs.python.org/3/tutorial/) and [documentation](https://docs.python.org/3/).

For grading, we will be using an autograder in [Gradescope](http://gradescope.com/). It grades based on values in specific variables or function names. Read the following instructions carefully. It is repeated for you again at the end of this notebook.

You should make sure any code that you write to answer the questions is included in this notebook. You are **required** to go to the Kernel option and choose **"Restart & Run All"** before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at Gradescope. **Submissions not prepared correctly as above will lose points.**

Make sure to also read the problem carefully to know what names to use so the autograder can correctly find your answers. 

There are 3 parts in this exercise, with first 2 required. The last part is optional (worth 0.0 points) but our autograder also gives feedback on it. You are therefore welcome to submit your answers to the optional part along with the required parts.

## Part 1: Central Tendency

This part is about [measures of central tendency](https://www.youtube.com/watch?v=MF6ij7tckI0&ab_channel=StudyHall). Watch the video in the link if you are unfamiliar with mean, median, and mode.

The goal of this section is to help you start building an intuition of the differences between the mean, median, and mode. It does this through brain teaser-like number manipulation problems.

### Question 1 (4 points)
Complete the function `central_measures` below that takes as input a list `lst` of (not necessarily distinct) numbers and an integer $k \geq 0$. It should output the mean, median, and mode of `lst` after the largest $k$ numbers in `lst` are removed. You may assume `lst` always contains at least $(k+1)$ numbers, so the resulting list is not empty. If there exist multiple modes, you may return any of them, or all of them as a list. 

You may find the built-in Python [`sorted`](https://docs.python.org/3/howto/sorting.html) function and the [`statistics`](https://docs.python.org/3/library/statistics.html#) module (imported as `stat` below for you) useful. You can also import and use other packages like `numpy` or `scipy` if you are comfortable with using them. Do not change the default variable names.

```python
import statistics as stat

# Code for question 1
def central_measures(lst, k): # DO NOT MODIFY THIS LINE
    mean = ...
    median = ...
    mode = ...
    return mean, median, mode # DO NOT MODIFY THIS LINE

# Your code should output (3, 2, 2) or (3, 2, [2]) for the example below
central_measures([1,2,2,4,6,6], 1)
```

### Question 2 (12 points, 4 each)
- Put in `q2_1` a list of 5 positive numbers such that its **mean** remains the same when one of its maximum values is removed (so it is only 4 long). 

- Put in `q2_2` a list of 5 positive numbers (of at least 3 distinct values) such that its **median** remains the same when one of its maximum values is removed. 

- Put in `q2_3` a list of 5 positive numbers such that its *unique* **mode** *changes* when its **two** largest values are removed. (Note: This means there should only be *one* mode both before and after the removal.)

You are welcome to use your `central_measures` function to help you figure out the answer to these questions.

```python
# Code for question 2
q2_1 = ...
q2_2 = ...
q2_3 = ...
```

### Question 3 (4 points)
Put in `q3` a list of 5 positive numbers such that:
- its mean is larger than 50;
- the absolute difference between its mean and median is at most 5;
- after removing one of its maximum value(s), the median does not change;
- after removing one of its maximum value(s), the absolute difference between the old mean and the new mean is larger than 25.

You are welcome to use your `central_measures` function to help you figure out the answer.

Hint: We recommend you start with a list containing the number 1, three numbers below 100, and then change the value of the maximum number(s) to see what happens.

```python
# Code for question 3
q3 = ...
```

## Part 2: Rides

The [boston.csv](boston.csv) file contains data on weekday Uber rides in the Boston, Massachusetts metropolitan areas from the [Uber Movement](https://movement.uber.com) project. The `sourceid` and `dstid` columns contain codes corresponding to start and end locations of each ride. The `hod` column contains codes corresponding to the hour of the day the ride took place. The `ride time` column contains the length of the ride, in minutes.

The code below will open the file and read the data as a list of rows, with each row represented as a string. It then prints the first four rows. Note that the first (index 0) row contains the column headers.

```python
# Do NOT change this code or the autograder may not work
f = open('boston.csv')
data = f.readlines()
f.close()
for row in data[:4]:
    print(row)
```

### Question 4 (4 points)
In the variable `q4` place how many rides are listed in the file.

```python
# Code for question 4
q4 = ...
```

### Question 5 (4 points)
In the variable `q5` place the length of the longest ride.

```python
# Code for question 5
q5 = ...
```

### Question 6 (4 points)
In the variable `q6` place the percent of rides that are under 10 minutes. The percent should be between 0 and 1.

```python
# Code for question 6
q6 = ...
```

### Question 7 (4 points)
For each HOD (hour of day), count how many rides take place during that hour. Put it in a list where index 0 has the number of rides where `hod` was 0, at index 1 where `hod` was 1, etc. until index 23. This means your list will be 24 long. Put this list in a variable called `q7`. Every element in `q7` should be an `int` type.

```python
# Code for question 7
q7 = ...
```

### Question 8 (4 points)
What are the three most common **start** locations (`sourceid`) for rides in the dataset? 

Place your answer in the variable `q8` sorted with the most common location first. **Every item in `q8` should be of a `str` type.** Although the `sourceid` look like numbers, they do not convey numerical meanings (e.g., taking the mean of them does not make sense), so they should be treated as strings. 

*Hint. You can also use `sorted` for this question. However, using `sorted` on a huge list will make your code very slow. Instead, think about using `sorted` on a smaller object (e.g., a dictionary).*

```python
# Code for question 8
q8 = ...
```

## Part 3: Text (Optional, 0 points, with autograder feedback)
The `text.txt` file contains the opening sentences of Charles' Dickens novel *A Tale of Two Cities*. Run the code below. It will read the file, store it in a string variable `text`, and print the first 500 characters.

```python
# Do NOT change this code or the autograder may not work
f = open('text.txt')
text = f.read()
f.close()

print(text[:500], '...')
```

### Question 9
In the variable `q9_1` put the first sentence of the text. In the variable `q9_2` put the last sentence of the text. You may assume sentences only end with a period (`'.'`).

You may find the Python `split` method for strings to be useful: https://docs.python.org/3/library/stdtypes.html?highlight=split#str.split `sep='.'`.

```python
# Code for question 9
q9_1 = ...
q9_2 = ...
```

### Question 10
In the variable `q10` place the total number of vowels (`a`, `e`, `i`, `o`, `u`) that appear in the text.

```python
# Code for question 10
q10 = ...
```

### Question 11
A trigram is a string of exactly 3 characters (or a string with length exactly 3), including blank spaces. For example, the first five trigrams in the text are (blank substituted by underline `_` for better demonstration):
1. `it_`
2. `t_w`
3. `_wa`
4. `was`
5. `as_`

In the variable `q11_1` place the most common trigram in the text. In the variable `q11_2` place the number of times the most common trigram shows in the text.

```python
# Code for question 11
q11_1 = ...
q11_2 = ...
```

<!-- BEGIN QUESTION -->

# AI Disclosure

Use the **Artificial Intelligence Disclosure (AID) Framework** to explain your use of AI on this assignment. Other headings you can use include:

- *Debugging:* Using AI to help you fix your code so that it works. You should state how you used it for this purpose.

Here are some examples:

*Artificial Intelligence Tools:* ChatGPT v5 via chatgpt.com. *Conceptualization:* I gave chatgpt.com the election data set and asked it for ideas on interesting statistics I could get from the data. *Methodology:* I asked it for help on how to write the code to get the statistic I chose, but I wrote the code myself. *Writing — Review & Editing:* I wrote out my explanation for what the statistic meant, then gave that text and the rubric to chatgpt and asked it to give me feedback on how to update the explanation to conform to the rubric.

*Artificial Intelligence Tools:* ChatGPT v4o via DukeGPT. *Information Collection:* DukeGPT was used to find the function needed to get the index value of the maximum value of a Series and the syntax needed to filter rows in a pandas dataframe using multiple columns. *Debugging:* DukeGPT was used to help me find a bug in my code for Q1 where I copied in the code and error, stated what the code should do, and asked for help.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **required** to go to the Kernel option and choose **"Restart & Run All"** before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at Gradescope. **Submissions not prepared correctly as above will lose points.**

Tests for check_bad_variable_names (points: [0.5])
```
>>> from auto_checks import get_default_type_definitions
>>> types_to_check = {'int': int, 'all': all, 'min': min, 'max': max, 'any': any, 'isinstance': isinstance, 'float': float, 'complex': complex, 'str': str, 'bytes': bytes, 'bytearray': bytearray, 'list': list, 'tuple': tuple, 'range': range, 'dict': dict, 'set': set, 'frozenset': frozenset, 'bool': bool, 'NoneType': type(None), 'Ellipsis': type(Ellipsis), 'NotImplemented': type(NotImplemented), 'slice': slice, 'memoryview': memoryview, 'classmethod': classmethod, 'staticmethod': staticmethod, 'property': property, 'object': object}
>>> def are_types_overwritten(types_to_check):
...     """
...     Joseph Scarpa, 02-24-2024
...     
...     Checks if built-in types are overwritten by custom definitions.
...     - types_to_check: A dictionary mapping from type names (as strings) to the custom types defined in the student notebook.
...     
...     Returns:
...     - Boolean indicating whether the built-in types are overridden or not.
...     """
...     built_in_types = get_default_type_definitions()
...     for type_name, type_val in types_to_check.items():
...         if built_in_types.get(type_name) is not type_val:
...             msg = f'NOT YET: It looks like you have created a variable called {type_name}. '
...             msg += 'This is a reserved keyword in Python. Use a different variable name.'
...             print(msg)
...             return False
...     return True
>>> are_types_overwritten(types_to_check)
True
```


Tests for q01 (points: [2.0, 2.0])
```
>>> from auto_checks import q1_format
>>> q1_format(central_measures)
'PASS'
```
```
>>> from auto_checks import q1_correctness
>>> q1_correctness(central_measures)
'PASS'
```


Tests for q02 (points: [2.0, 1.0, 1.0, 2.0, 2.0, 2.0, 2.0])
```
>>> from auto_checks import q2_format
>>> q2_format(q2_1, 'q2_1')
'PASS'
```
```
>>> from auto_checks import q2_format
>>> q2_format(q2_2, 'q2_2')
'PASS'
```
```
>>> from auto_checks import q2_2_format
>>> q2_2_format(q2_2)
'PASS'
```
```
>>> from auto_checks import q2_format
>>> q2_format(q2_3, 'q2_3')
'PASS'
```
```
>>> from auto_checks import q2_1_correctness
>>> q2_1_correctness(q2_1)
'PASS'
```
```
>>> from auto_checks import q2_2_correctness
>>> q2_2_correctness(q2_2)
'PASS'
```
```
>>> from auto_checks import q2_3_correctness
>>> q2_3_correctness(q2_3)
'PASS'
```


Tests for q03 (points: [0, 1.0, 1.0, 1.0, 1.0])
```
>>> from auto_checks import q2_format
>>> q2_format(q3, 'q3')
'PASS'
```
```
>>> from auto_checks import q3_correctness_1
>>> q3_correctness_1(q3)
'PASS'
```
```
>>> from auto_checks import q3_correctness_2
>>> q3_correctness_2(q3)
'PASS'
```
```
>>> from auto_checks import q3_correctness_3
>>> q3_correctness_3(q3)
'PASS'
```
```
>>> from auto_checks import q3_correctness_4
>>> q3_correctness_4(q3)
'PASS'
```


Tests for q04 (points: [2.0, 2.0])
```
>>> isinstance(q4, int)
True
```
```
>>> from auto_checks import q4_value
>>> q4_value(q4)
'PASS'
```


Tests for q05 (points: [2.0, 2.0])
```
>>> isinstance(q5, float)
True
```
```
>>> from auto_checks import q5_value
>>> q5_value(q5)
'PASS'
```


Tests for q06 (points: [2.0, 2.0])
```
>>> isinstance(q6, float)
True
```
```
>>> from auto_checks import q6_value
>>> q6_value(q6)
'PASS'
```


Tests for q07 (points: [1.0, 1.0, 2.0])
```
>>> isinstance(q7, list)
True
```
```
>>> all((isinstance(item, int) for item in q7))
True
```
```
>>> from auto_checks import q7_value
>>> q7_value(q7)
'PASS'
```


Tests for q08 (points: [0.5, 0.5, 1.0, 2.0])
```
>>> isinstance(q8, list)
True
```
```
>>> len(q8) == 3
True
```
```
>>> from auto_checks import q8_type
>>> q8_type(q8)
'PASS'
```
```
>>> from auto_checks import q8_value
>>> q8_value(q8)
'PASS'
```


Tests for q09 (points: [0, 0, 0, 0, 0, 0])
```
>>> isinstance(q9_1, str)
True
```
```
>>> isinstance(q9_2, str)
True
```
```
>>> len(q9_1) == 611
True
```
```
>>> len(q9_2) == 204
True
```
```
>>> from auto_checks import q9_1_value
>>> q9_1_value(q9_1)
'PASS'
```
```
>>> from auto_checks import q9_2_value
>>> q9_2_value(q9_2)
'PASS'
```


Tests for q10 (points: [0, 0])
```
>>> isinstance(q10, int)
True
```
```
>>> from auto_checks import q10_value
>>> q10_value(q10)
'PASS'
```


Tests for q11 (points: [0, 0, 0, 0])
```
>>> isinstance(q11_1, str)
True
```
```
>>> isinstance(q11_2, int)
True
```
```
>>> from auto_checks import q11_1_value
>>> q11_1_value(q11_1)
'PASS'
```
```
>>> from auto_checks import q11_2_value
>>> q11_2_value(q11_2)
'PASS'
```
