# Practicum 1

By submitting this notebook to Gradescope you are acknowledging that you are upholding the Duke honor code that this notebook is representative of the work of all that are on the submission and no one else.

This assessment can be done in teams of two. It represents your pair's work, and only your pair's work to show your all's understanding of the material. This document and your answers are for your pair's eyes only.

This exam is open book, open note, open Internet, open to large language models, and closed to people outside of your team. This means you cannot communicate with someone outside of your pair about the practicum until grades are published, including asking someone through the Internet (like stackoverflow) to receive help.

If you wish to earn an Exemplary (full credit), your plots must be clean and easy to understand. That means ensuring the text of axis labels are not overlapping by using parameters like `aspect` and `height`. Using column names as your axis labels is fine and titles are not required. **In addition, all answers that are clearly plagiarized from the practice exam will be marked as Not Yet**. 

You **may reuse** the results of prior questions (like variables or new columns) in later questions. If you do so, you need to clearly state it for the graders.

Except for the Monte Carlo Simulation problem (question 4), you may **not use** `for` or `while` loops if there are clear ways to do it using libraries and without such loops. In addition, the code must guarantee that the loop will iterate less than 25 times, including if we swap out the data for another data file with 10x more rows. Using such loops will immediately mark your solution as "Not Yet" because one large purpose of this exam is to assess your skills at taking advantage of the data science libraries.

If you think something is ambiguous or some assumption/key information is missing, decide on your own interpretation or any additional assumption and state them clearly.

The grading for this exam will be very lenient in the reasonableness of how you use the data since we do not expect you to have a great understanding of the data set. Moreover, the values in this data set do not necessarily reflect the real world.

# Practicum Update Diff Cell (2 points, if you submit the Update)
If you are not updating your practicum for the Practicum Update, ignore this cell. If you are updating this notebook for the Practicum Update, fill out this template to make clear what you did and did not change.

### Question 1
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question, for example something like "We did not change the code and only changed the writeup to better answer questions 2 and 3.")

### Question 2
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question, for example something like "We cleaned up the code by removing the `==True` and did not change the writeup.")

### Question 3
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question)

### Question 4
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question)

# Overview Note

We recommend scanning through the entire exam before beginning so you have a sense of what each question is asking you to do.

Feel free to add columns to the dataframe and use them throughout the notebook. Just note in your writeup where the code for a particular new column is to make grading easier.

**Using LLMs**: For each question, there is a variable called `q#_llm_flag`. Set this to `True` if you used an LLM to help you with that question and `False` otherwise.

```python
# Standard imports
import numpy as np
import pandas as pd
import re
import seaborn as sns
import matplotlib.pyplot as plt
```

# Group Members

In the below cell provide the NetIds as strings of all the people working on this practicum. If you are working solo, write `student2`'s NetId as `'no1'`.

```python
student1 = 'mz248'
student2 = 'no1'
```

# Data Set: AirBnB Listings

Below is a data set about AirBnB listings in North Carolina. Every row represents a single AirBnB listing.

The meanings of all columns are listed below.

Note: The values in this data set do not necessarily reflect the real world and may be inaccurate.

| Variable | Type | Description |
|----------|------|-------------|
| `id` | int64 | ID for the listing |
| `name` | str | Name of the listing |
| `description` | str | Description of the listing |
| `neighborhood_overview` | str | Overview of the neighborhood the listing is in |
| `host_id` | int64 | ID for the host |
| `host_name` | str | Host's name, could be one or more people, or a company |
| `host_since` | str | Date since being a host |
| `host_location` | str | Where the host is from with format "City, State", state is the state abbreviation |
| `host_about` | str | Description about the host |
| `host_response_time` | str | How quickly the host responds; 4 possible options |
| `host_response_rate` | str | Percent of times host responds as a string |
| `host_acceptance_rate` | str | Percent of times host accepts a booking as a string |
| `host_is_superhost` | bool | Whether the host is a super host |
| `host_listings_count` | int64 | Number of listings the host has in NC |
| `host_total_listings_count` | int64 | Number of listings the host has total |
| `host_has_profile_pic` | bool | Whether the host has a picture in their profile |
| `host_identity_verified` | bool | Whether the host's identity has been verified |
| `neighbourhood` | str | Where the listing is with format "City, State, Country", state is the full state name |
| `neighbourhood_zip` | int64 | The neighborhood's zip code |
| `latitude` | float | Listings latitude |
| `longitude` | float | Listings longitude |
| `property_type` | str | Type of property, not standardized; there are many possible options |
| `room_type` | str | Type of room, 4 possible options |
| `accommodates` | int64 | How many people the listing can have |
| `bathrooms_text` | str | Text description of the bathrooms |
| `bedrooms` | float | Number of bedrooms |
| `beds` | float | Number of beds |
| `amenities` | str | Python formatted list of amenities with no particular order |
| `price` | str | Price of one night as a string |
| `minimum_nights` | int64 | Minimum number of nights for a booking |
| `maximum_nights` | int64 | Maximum number of nights for a booking |
| `availability_365` | int64 | Number of days available throughout the year |
| `review_scores_rating` | float | Mean general rating |
| `review_scores_accuracy` | float | Mean rating about the listing's accuracy |
| `review_scores_cleanliness` | float | Mean rating about the listing's cleanliness |
| `review_scores_checkin` | float | Mean rating about the listing's checkin experience |
| `review_scores_communication` | float | Mean rating about the listing's communication with the host |
| `review_scores_location` | float | Mean rating about the listing's location |
| `review_scores_value` | float | Mean rating about the listing's value |
| `instant_bookable` | bool | Whether the listing is instantly bookable |
| `calculated_host_listings_count` | int64 | A different calculated count of the number of listings by the host |
| `calculated_host_listings_count_entire_homes` | int64 | How many entire homes the host lists |
| `calculated_host_listings_count_private_rooms` | int64 | How many private rooms the host lists |
| `calculated_host_listings_count_shared_rooms` | int64 | How many shared rooms the host lists |

```python
# import data
df = pd.read_csv('listings.csv')
print(df.shape)
df.head()
```

Output:
```
(3161, 44)
       id                                               name  \
0  108061  Walk to stores/parks/downtown. Fenced yard/Pet...   
1  155305                 Cottage! BonPaul + Sharky's Hostel   
2  156805                  Private Room "Ader" at BPS Hostel   
3  197263                       Tranquil Room & Private Bath   
4  209068                                    Terrace Cottage   

                                         description  \
0  Walk to town in ten minutes! Monthly rental in...   
1  <b>The space</b><br />Pet friendly private cot...   
2  <b>The space</b><br />Private Rooms at Bon Pau...   
3  This is a comfy, peaceful and clean room with ...   
4  Located in one of Asheville's oldest historic ...   

                               neighborhood_overview  host_id host_name  \
0  I love my neighborhood! Its friendly, easy-goi...   320564      Lisa   
1  We are within easy walk of pubs, breweries, mu...   746673   BonPaul   
2  Easy walk to pubs, cafes, bakery, brewe
```

# Question 1 (24 points: 4 autograded, 20 manually graded)

Do the following:
1. Create a mask that you will use to filter the data:
    1. It must be True for 10%-90% of the data and False otherwise.
    1. You may base the mask on multiple columns, but you do not have to.
    1. Add the mask as a new column to the dataframe.
    1. Put the new column's name in `new_mask_col_name`.
1. Put the fraction of rows kept by the mask in `frac_data_kept`.
1. Use the mask to filter the data and then calculate a single descriptive statistic (e.g., mean, count, median, proportion, etc.) of a **different** column. Put this value in `statistic_mask`.
1. Calculate the same descriptive statistic for the entire data set.  Put this value in `statistic_all`.

**See the written portion of this question to see what you must explain.**

## Answer 1: Code

```python
q1_llm_flag = True
new_mask_col_name = 'price_middle_80pct'
frac_data_kept = pd.to_numeric(df['price'].str.replace(r'[\$,]', '', regex=True), errors='coerce').between(pd.to_numeric(df['price'].str.replace(r'[\$,]', '', regex=True), errors='coerce').quantile(0.1), pd.to_numeric(df['price'].str.replace(r'[\$,]', '', regex=True), errors='coerce').quantile(0.9)).mean()
df[new_mask_col_name] = pd.to_numeric(df['price'].str.replace(r'[\$,]', '', regex=True), errors='coerce').between(pd.to_numeric(df['price'].str.replace(r'[\$,]', '', regex=True), errors='coerce').quantile(0.1), pd.to_numeric(df['price'].str.replace(r'[\$,]', '', regex=True), errors='coerce').quantile(0.9))
statistic_mask = df[df[new_mask_col_name]]['minimum_nights'].mean()
statistic_all = df['minimum_nights'].mean()

# Leave these here to help with grading.
print("New mask's column name:", new_mask_col_name)
print('Fraction of data kept:', frac_data_kept)
print('Descriptive statistic of masked data:', statistic_mask)
print('Descriptive statistic of all data:', statistic_all)
```

Output:
```
New mask's column name: price_middle_80pct
Fraction of data kept: 0.8041758937045239
Descriptive statistic of masked data: 5.0098347757671124
Descriptive statistic of all data: 5.197722239797533
```

<!-- BEGIN QUESTION -->

## Answer 1: Written portion
Do the following:
1. Explain what your mask is filtering for.
1. State the values of and what your descriptive statistics in `statistic_mask` and `statistic_all` are. Explain what these numbers tell you about the data. Are their values surprising? Why or why not?
1. If you did not in the prior subquestions, compare the two values, is one bigger? Which one? Speculate why this is the case.

ANSWER
1. My mask filters for listings where the price falls between the 10th and 90th percentiles, which means I am keeping the middle 80% of the price distribution and removing the most extreme low and high prices. This is useful for focusing on the core pricing data and ignoring outlier listings that might skew the analysis.

2. The masked group's average minimum nights is shown in `statistic_mask`, while `statistic_all` gives the overall average for the full data. Comparing these numbers helps us see if the listings with typical prices have different minimum night requirements than the data as a whole. In my analysis, the numbers are close, which suggests the outlier-priced listings do not have a drastically different policy for minimum stays.

3. The statistic for the masked data is slightly lower than for all listings, which means the listings in the price middle 80% tend to allow shorter stays on average. This could make sense because extremely inexpensive or expensive listings might have special requirements, like longer minimum stays to offset lower revenue or to match a premium offering. This comparison gives us insight into how pricing and stay requirements are related in the dataset.

<!-- END QUESTION -->

# Question 2 (25 points: 4 autograded, 21 manually graded)

Assessment goal: Demonstrate an understanding of how to use regular expressions to extract information from strings that are irregular.

Do the following:
1. Use a regular expression(s) to create a new column by extracting information (a substring or number) from an existing column with data of type string.
    1. The existing/original column must have more than 18 unique values. So, do not use something like `room_type`.
    1. The extracted information cannot be from something that is extremely uniform, such as something that is of fixed width and in the same location, like removing "North Carolina" from `neighbourhood` or the `$` from `price`. If you could get it by just slicing all the strings the same way, it is not a good demonstration of using regular expressions.
    1. The extracted information must be information that cannot easily be extracted with methods other than regular expressions (such as creating a `True`/`False` or `"yes"`/`"no"` column by just using `str.contains` or equivalent).
    1. Therefore, the extracted information must be a number or of type string.
    1. Put the name of the column you are extracting from in `org_col_name`.
    1. Put the new column's name in `new_col_name`.
1. Count the number of NaNs in **the original column** and put it in `nan_count_org`.
1. Count the number of NaNs in **your new column** and put it in `nan_count_new`.

**See the written portion of this question to see what you must explain.**

## Answer 2: Code

```python
q2_llm_flag = True
org_col_name = 'name'
new_col_name = 'num_in_name'
df[new_col_name] = df[org_col_name].str.extract(r'(\d+)')
nan_count_org = df[org_col_name].isna().sum()
nan_count_new = df[new_col_name].isna().sum()


# Leave these here to help with grading.
print('Original column:', org_col_name)
print('Unique values in original column:', df[org_col_name].nunique())
print('New column:', new_col_name)
print('Head of old column:')
print(df[org_col_name].head(5))
print('Head of the non-NaN values of the new column:')
print(df[new_col_name].dropna().head(5))
print('Count of NaNs in original column:', nan_count_org)
print('Count of NaNs in new column:', nan_count_new)
print('% NaNs in new column:', 100*nan_count_new/df.shape[0])
```

Output:
```
Original column: name
Unique values in original column: 3143
New column: num_in_name
Head of old column:
0    Walk to stores/parks/downtown. Fenced yard/Pet...
1                   Cottage! BonPaul + Sharky's Hostel
2                    Private Room "Ader" at BPS Hostel
3                         Tranquil Room & Private Bath
4                                      Terrace Cottage
Name: name, dtype: object
Head of the non-NaN values of the new column:
10        8
12        2
14       90
25    28803
36       10
Name: num_in_name, dtype: object
Count of NaNs in original column: 0
Count of NaNs in new column: 2278
% NaNs in new column: 72.06580196140462
```

<!-- BEGIN QUESTION -->

## Answer 2: Written portion
Do the following:
1. Explain what the code you wrote in the code cell does.
1. State how many NaNs there are in the original and new columns. Provide a possible reason why the NaNs are/are not there.
1. If you did not discuss this in the prior subquestion, why do the original and new columns have the same or not the same number of NaNs?

ANSWER 

The code first creates a new column called num_in_name by extracting any sequence of digits from the name column using a regular expression. This allows us to capture numerical information embedded in listing names, which could be useful for further analysis.

The number of NaN values in the original name column is 0, which means every row originally had a value in that column. The new column num_in_name has more NaNs than the original, because only some names actually contain numbers for the regular expression to extract, so it leaves a NaN where no match is found.

The difference in NaN counts between the columns happens because the extraction with str.extract only returns a value when digits are present in the name otherwise it puts a NaN. By contrast, the original name column only has NaNs if the name itself is missing.


<!-- END QUESTION -->

# Question 3 (25 points: 4 autograded, 21 manually graded)

For this question, do the following:

1. `col_cat` column - Create a new column that labels each row with some kind of category. Put the name of your new column in `col_cat`. Other requirements include:
   1. The column needs at least 2 categories and cannot have more than 18.
   1. You may use the results of the regular expression question, just state as much in the write up.
   1. If you wish, you may filter to reduce the number of categories or the data you are visualizing. For example, you can filter to make your plot more legible. You must have at least 100 rows in the dataframe you use to create your plot.
1. Put in `val_counts` a Series where the key/index are the categories of `col_cat` and the values are the number of rows with that category.
1. Use `relplot`, `displot`, or `catplot` to create a plot using the `col_cat` column and **at least one other column** ([documentation figure](https://seaborn.pydata.org/tutorial/function_overview.html#figure-level-vs-axes-level-functions "‌")).
1. Put the data you used for the plot in `table3`. This table should include **only** the data you used and no other columns nor rows.

**Note**: Remember, your plot must be legible to be able to earn an Exemplary mark.

**See the written portion of this question to see what you must explain.**

## Answer 3: Code

```python
q3_llm_flag = ...

col_cat = ...
table3 = ...

val_counts = ...

# There should be at least one call to a seaborn function
# sns.?


# Leave these here to help with grading.
print('Shape of dataframe:', table3.shape)
print('At least 100 rows in data?:', table3.shape[0] >= 100)
print('Columns:', table3.columns)
print('Number of unique values in', col_cat, 'column:', table3[col_cat].nunique())
print('Number of categories within range?', 2 <= table3[col_cat].nunique())
print('Counts of', col_cat, 'categories:')
print(val_counts[:15])
table3.head()
```

<!-- BEGIN QUESTION -->

## Answer 3: Written portion
Explain the following:
1. What is the code doing? Include:
    1. What is `col_cat`? What do the categories in it mean? How did you create it?
    1. How you filtered and why you chose to keep the data you did. If you did not filter at all, just state as much.
1. Explain generally what the plot is showing.
1. Pick a single spot on the plot and describe it by filling out the sentence "The \_\_ on the chart means \_\_ ." (you may use these exact words)
    1. Make sure it is clear what spot you are referring to and you are "pointing" it out without using its data values (otherwise, that's like defining a word while using the word in the definition). For example, "the top right blue dot" or "the shortest red bar," as opposed to "the top dot for AirBnBs that are Private rooms."
    1. If using something like a histogram, be precise about what is inclusive vs. exclusive. For example, “it includes up to 10” is not the same as “it includes up to and including 10.”
    1. If using a boxplot, make sure to mention multiple lines/parts of the boxes and note any outliers.
1. Why did you choose `relplot`/`displot`/`catplot` and use any deviations from the default behavior of that function? Discuss this at least partially in terms of the insight you sought from the plot.
1. Why the other `seaborn` functions would have been less useful to show you what you wanted to see OR why they are equally viable?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Question 4 (25 points: 4 autograded, 21 manually graded)

In this question, you will run a Monte Carlo Simulation. Your code will sample a subset of listings. Describe a scenario and then calculate the probability of that scenario happening with the following constraints:
1. Use at least 2 columns of data as part of your scenario.
1. The scenario should have some kind of inequality in it, such as "The mean cost of AirBnBs in Fletcher is more than Asheville."
1. The scenario should involve more than one AirBnB.
1. You may filter the data if you wish, but there must be at least 100 rows in the data set you simulate over. The column(s) used in the filtering can count towards your "at least 2 columns" constraint. If you filter by a random sample, you must explain why in the written portion of this question.
1. Run enough simulations that you are confident in your answer within 0.1% (this code prints things as a decimal, so 0.1% is the third digit after the decimal). Do not change the `count` and `sims` variable names. The autograder uses these to do its sanity check.

**See the written portion of this question to see what you must explain.**

**Note**: You are allowed to use a `for` or `while` loop for this question.

## Answer 4: Code

```python
q4_llm_flag = True

# Scenario: Probability that a randomly selected Asheville Airbnb is more expensive than a non-Asheville Airbnb.
# Preprocess price column only once
prices = pd.to_numeric(df['price'].replace('[\$,]', '', regex=True), errors='coerce')

# Filter once for efficiency
loc = df['host_location'].fillna('')
ash_mask = loc.str.lower().str.contains('asheville')
notash_mask = ~ash_mask

ash_prices = prices[ash_mask].dropna().values
notash_prices = prices[notash_mask].dropna().values

assert len(ash_prices) >= 50 and len(notash_prices) >= 50, "Need at least 50 rows in each pool"

sims = 100_000  # Large N for confidence
# Vectorized sampling for efficiency
ash_samples = np.random.choice(ash_prices, sims)
notash_samples = np.random.choice(notash_prices, sims)

count = np.sum(ash_samples > notash_samples)

for i in range(1000, sims+1, 1000):
    if i <= sims:
        print(i, np.sum(ash_samples[:i] > notash_samples[:i]) / i)

# Leave this here to help with grading.
print('Final result:', count/sims)
```

Output:
```
1000 0.422
2000 0.4355
3000 0.442
4000 0.44675
5000 0.4508
6000 0.448
7000 0.44642857142857145
8000 0.44575
9000 0.4491111111111111
10000 0.4522
11000 0.44945454545454544
12000 0.44866666666666666
13000 0.44746153846153847
14000 0.448
15000 0.44766666666666666
16000 0.447625
17000 0.44882352941176473
18000 0.4487777777777778
19000 0.4475789473684211
20000 0.4478
21000 0.4471904761904762
22000 0.44777272727272727
23000 0.4474347826086956
24000 0.446625
25000 0.44572
26000 0.4446538461538462
27000 0.4457037037037037
28000 0.44589285714285715
29000 0.44624137931034485
30000 0.4463
31000 0.44587096774193546
32000 0.44528125
33000 0.445
34000 0.4444705882352941
35000 0.44454285714285713
36000 0.4446111111111111
37000 0.4444054054054054
38000 0.4441842105263158
39000 0.443974358974359
40000 0.44415
41000 0.44385365853658537
42000 0.44357142857142856
43000 0.4434186046511628
44000 0.4435909090909091
45000 0.44382222222222223
46000 0.44382608695652176
47000 0.44363829787234044
48000 0.44379166
```

<!-- BEGIN QUESTION -->

## Answer 4: Written portion
Do the following:
1. Provide a 1-3 sentence explanation of the scenario. Be explicit in whether this scenario requires with or without replacement and why.
1. State the probability your simulation found, rounded to the nearest 0.1%, and explain why you are confident in your answer.

1. In this scenario, we estimate the probability that a randomly selected Airbnb listing with `host_location` containing Asheville has a higher nightly price than a listing whose `host_location` does not contain Asheville. This simulation is with replacement because each draw is independent across trials, so the same listing can be selected again in different simulations.
2. My simulation found a probability of about 44.4% (final estimate `0.44353`, rounded to the nearest 0.1%). I am confident in this estimate because it is based on 100,000 simulations, and the running estimate stabilizes to the third decimal place, which supports accuracy at the required 0.1% level.

<!-- END QUESTION -->

# AI Disclosure (1 point)
Use the [Artificial Intelligence Disclosure (AID) Framework](https://crln.acrl.org/index.php/crlnews/article/view/26548/34482) to explain your use of AI on this exam. Keep in mind the other examples on the class's policy web page.

## My AI Disclosure

*Artificial Intelligence Tools*: 1 2 3 replace me!

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **strongly encouraged** to go to the Kernel option and choose **"Restart & Run All"**  before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/ "‌"). **Submissions not prepared correctly as above may lose points.**

Tests for q0-sanity-check (points: [0, 0])
```
>>> student1.lower() != 'no1' and re.fullmatch('[a-z]+[0-9]*', student1.lower()) != None
True
```
```
>>> re.fullmatch('[a-z]+[0-9]*', student2.lower()) != None
True
```


Tests for q1-code (points: [0, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5, 1])
```
>>> isinstance(q1_llm_flag, bool)
True
```
```
>>> isinstance(new_mask_col_name, str)
True
```
```
>>> new_mask_col_name in df.columns
True
```
```
>>> bool(pd.api.types.is_bool_dtype(df[new_mask_col_name]))
True
```
```
>>> isinstance(frac_data_kept, float) or isinstance(frac_data_kept, np.float64)
True
```
```
>>> bool(0.1 <= frac_data_kept and frac_data_kept <= 0.9)
True
```
```
>>> from numbers import Number
>>> isinstance(statistic_mask, Number)
True
```
```
>>> from numbers import Number
>>> isinstance(statistic_all, Number)
True
```


Tests for q2-code (points: [0, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5, 0.5])
```
>>> isinstance(q2_llm_flag, bool)
True
```
```
>>> isinstance(org_col_name, str)
True
```
```
>>> bool(df[org_col_name].nunique() > 18)
True
```
```
>>> isinstance(new_col_name, str)
True
```
```
>>> new_col_name in df.columns
True
```
```
>>> isinstance(nan_count_org, int) or isinstance(nan_count_org, np.integer)
True
```
```
>>> bool(nan_count_org >= 0)
True
```
```
>>> isinstance(nan_count_new, int) or isinstance(nan_count_new, np.integer)
True
```
```
>>> bool(nan_count_new >= 0)
True
```


Tests for q3-code (points: [0, 1, 1, 1, 1])
```
>>> isinstance(q3_llm_flag, bool)
True
```
```
>>> isinstance(col_cat, str)
True
```
```
>>> bool(table3[col_cat].nunique() >= 2 and table3[col_cat].nunique() <= 18)
True
```
```
>>> bool(table3.shape[0] >= 100)
True
```
```
>>> isinstance(val_counts, pd.Series)
True
```


Tests for q4-code (points: [0, 0.5, 1, 0.5, 1, 1])
```
>>> isinstance(q4_llm_flag, bool)
True
```
```
>>> isinstance(count, int) or isinstance(count, np.int64)
True
```
```
>>> bool(count > 0)
True
```
```
>>> isinstance(sims, int) or isinstance(sims, np.int64)
True
```
```
>>> bool(sims > 10)
True
```
```
>>> bool(count <= sims)
True
```
