# Homework 08

In this exercise, you will demonstrate your learning of prediction, including regression and classification, using the [scikit-learn library](https://scikit-learn.org/stable/index.html). Note that scikit-learn is included in the Anaconda distribution, so you should not need to install anything new to use it.

Show your work (code or otherwise) and/or briefly explain your answers. In general you will not receive full credit for numeric answers with no accompanying work or justification (math, code, explanation). For numeric answers, we will accept answers that are very slightly off due to rounding.

```python
# Run this code cell to import relevant libraries
# Please update to seaborn v0.11 or above for this assignment
# check version: print(sns.__version__)

import numpy as np
import pandas as pd
import seaborn as sns
sns.set_theme()
```

## Update scikit-learn library

Run the following cell to install or update the `Scikit-learn` library if not already done. Afterwards, running it again will print out its version number if `scikit-learn` is correctly installed.

```python
try:
    import sklearn
    print(sklearn.__version__)
    if (sklearn.__version__ < '1.8.0'): #update for latest scikit-learn version
        %pip install -U --user scikit-learn
except Exception as e:
    %pip install scikit-learn
```

### Question 1.1 (6 points; 4 autograded)
Below we import the `tips` dataset containing information from several meals in a particular restaurant. We preview the first few rows.

```python
# Run but do not modify this code

tips = sns.load_dataset("tips")
print(tips.shape)
tips.head()
```

We are interested in modeling `tip` as a function of `total_bill` and `size`. Use scikit-learn to perform linear regression to model `tip` as a function of `total_bill` and `size`. Once you have fit the model, obtain the model predictions and compare them to the real `tip`. Compute both the MSE (mean squared error) and the $r^2$ score (the square of Pearson's correlation coefficient). Report the MSE in `q1_mse` and the $r^2$ score in `q1_r2`, both are float/np.float.

To contextualize the performance of your model, consider a baseline model which simply predicts the median `tip` for every transaction. Compute the MSE and $r^2$ scores for this baseline, and report them in `q1b_mse` and `q1b_r2`.

Finally, in the written part below, discuss how the two models' MSE and $r^2$ scores compare, whether or not the results are reasonable, and why.

Notes:
- We import the relevant code from `sklearn` for you; you can find the documentation of the functions by searching the [scikit-learn library](https://scikit-learn.org/stable/index.html).
- You do not need to split the data into train/test sets for this problem, you can just use the whole dataset.
- You can use the default LinearRegression model, i.e., you do not need to pass any parameters to the constructor.
- You may find [`np.full`](https://numpy.org/doc/stable/reference/generated/numpy.full.html) or [`np.full_like`](https://numpy.org/doc/stable/reference/generated/numpy.full_like.html) useful.

```python
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
from sklearn.metrics import r2_score

# Write your code for question 1.1 here
# and store your results in the correct variable names for full credit

q1_mse = ... # float: MSE of your model
q1_r2 = ... # float: R^2 score of your model
q1b_mse = ... # float: MSE of the baseline 
q1b_r2 = ... # float: R^2 score of the baseline


# Leave these lines here to print the result
print("Model MSE:", q1_mse, "r^2:", q1_r2)
print("Baseline MSE:", q1b_mse, "r^2:", q1b_r2)
```

<!-- BEGIN QUESTION -->

### Answer 1.1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 1.2 (7 points; 4 autograded)
Now let’s incorporate `day` into our model. However, even though day is ordinal it is not clear whether the order should be salient in the model or not. Therefore we will build two models and compare them. Your first model should encode `day` using the ordinal encoder. You must specify the order for an OrdinalEncoder or it will use whatever order is in the dataframe. In this case, let's specify the `categories` parameter as `['Thur', 'Fri', 'Sat', 'Sun']`. The second model should use one hot encoding. Report the MSEs in `q12_1_mse` and `q12_2_mse` and the $r^2$ score in `q12_1_r2` and `q12_2_r2`, all are float/np.float. 

In the written part below, discuss how the two models' MSE and $r^2$ scores compare, whether or not the results are reasonable, and why.

Notes:
- We import the relevant code from `sklearn` for you; you can find the documentation of the functions by searching the [scikit-learn library](https://scikit-learn.org/stable/index.html).
- You do not need to split the data into train/test sets for this problem, you can just use the whole dataset.
- You can use the default LinearRegression model, i.e., you do not need to pass any parameters to the constructor.

```python
from sklearn.preprocessing import OrdinalEncoder, OneHotEncoder

# Write your code for question 1.2 here
# and store your results in the correct variable names for full credit

q12_1_mse = ... # float: MSE of your model using ordinal encoder
q12_1_r2 = ... # float: R^2 score of your model using ordinal encoder
q12_2_mse = ... # float: MSE of your model using one hot encoder
q12_2_r2 = ... # float: R^2 score of your model using one hot encoder


# Leave these lines here to print the result
print("Ordinal Encoding MSE:", q12_1_mse, "r^2:", q12_1_r2)
print("One Hot Encoding MSE:", q12_2_mse, "r^2:", q12_2_r2)
```

<!-- BEGIN QUESTION -->

### Answer 1.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2 (7 points; 3 autograded)
Below we import and preview the `penguins` dataset. We have previously modeled the numerical `flipper_length_mm` feature. In this question, we will model the categorical `species` feature.

```python
# Run but do not modify this code
# Only 11 penguins out of over 300 are missing any data;
# we simply drop those here
penguins = sns.load_dataset("penguins").dropna()
print(penguins.shape)
penguins.head()
```

To begin, we hold out 70% of the data randomly as a test set.

```python
# Run but do not modify this code

from sklearn.model_selection import train_test_split

q2_data = penguins.drop(columns=["species"])
q2_target = penguins["species"]
q2_train_data, q2_test_data, q2_train_target, q2_test_target = train_test_split(
    q2_data, q2_target, test_size=0.7, random_state=999)
```

Now, create and fit a logistic regression model to predict the `species` of a penguin using all of the other features in the dataset.  Note that two of the features, `island` and `sex`, are categorical. Use one hot encoding to encode these categorical features in your model and combine them with the numerical features. You should fit your model only on the training data, i.e., `q2_train_data` and `q2_train_target`.

Once you have fit your logistic regression model, make predictions for `q2_test_data` and then compute and report the `accuracy_score` of your predictions versus `q2_test_target`; put the `accuracy_score` in `q2_score` as `float`/`numpy.float64`. Also plot the confusion matrix for your predictions; see worked examples for more instructions on plotting confusion matricies. 

If everything goes right, your model should correctly distinguish a `species` from the other two (i.e., when your model predicts this species in the test data, it is always right).

The performace of your model is really good, but we do not know which feature(s) contributed the most. To find out if the categorical features help, fit a simpler version of your logistic regression model that uses only the numerical features (i.e., everything except `island` and `sex`) in the dataset as a baseline. Repeat the same process, put its `accuracy_score` in `q2b_score`, and plot its confusion matrix. 

In the written part, discuss the following:
1. How the two models' accuracy_score and their confusion matrices compare, whether or not the results are reasonable, and why.
2. Which specie(s) were correctly predicted every time with the data set and how you know this by stating specific numbers from your confusion matrices.

You are also welcomed to compare with a model that only uses the categorical features yourself, but that is not required in this homework.

Details:
- We import the relevant code from `sklearn` for you; you can find the documentation of the functions by searching the [scikit-learn library](https://scikit-learn.org/stable/index.html).
- If you use the default LogisticRegression model, you will likely get a warning message about convergence during fitting. We recommend you set the parameter `max_iter=10000` when calling the constructor (i.e., `LogisticRegression(max_iter=10000)` for both models. This will allow the optimization code for the model to converge most of the time. If you still have a warning, increase it more.
- If your confusion matrix has distracting white lines, add a line `plt.grid(False)` after each call to `ConfusionMatrixDisplay` to remove them.

```python
from sklearn.linear_model import LogisticRegression
from sklearn.preprocessing import OneHotEncoder
from sklearn.metrics import accuracy_score
from sklearn.metrics import ConfusionMatrixDisplay
from sklearn import metrics
import matplotlib.pyplot as plt # to better format confusion matrix

# Write your code for question 2 here
q2_score = ... # float: accuracy score of the prediction model
q2b_score = ... # float: accuracy score of the baseline model


# Leave these lines here to print the result
print("Model classification accuracy:", q2_score)
print("Baseline classification accuracy:", q2b_score)
```

<!-- BEGIN QUESTION -->

### Answer 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 3 (7 points, 4 autograded)
An interesting observation about classification in the bag of words model is that the predictive model knows nothing about the semantics of the particular language. For this question, you will build a sentiment analysis model of casual social media messages in the Central Thai language. The data consist of raw text messages, each of which is labeled with a **sentiment**, either negative `neg`, neutral `neu`, positive `pos`, or question `q`. For example, a message that reads "I am happy you are my friend" might be labeled as positive whereas one that reads "It is raining right now" might be labeled as neutral. The dataset is already split into a training and test set. We first import and preview the first ten messages and their sentiment labels.

```python
# Run but do not modify this code

# Training data
q3_train_text = pd.read_table("train_text.txt", dtype=str).squeeze("columns")
q3_train_label = pd.read_table("train_label.txt", dtype=str).squeeze("columns")

# # Held testing data
q3_test_text = pd.read_table("test_text.txt", dtype=str).squeeze("columns")
q3_test_label = pd.read_table("test_label.txt", dtype=str).squeeze("columns")

print("Training messages:", q3_train_text[:10].values, "\n")
print("Training message labels/classes:", q3_train_label[:10].values)
```

Now, create and fit a Naive Bayes model that predicts the sentiment label of a message given the raw text of the message. To do so, you will need to vectorize the messages using the `CountVectorizer` and then create a `MultinomialNB` model. Both are imported for you below; you can use default values for both. You are welcome but not required to combine the steps and use a pipeline object. Fit your model on the **training data** and then make predictions of the sentiment labels for `q3_test_text`. Compute and report the `accuracy_score` of your predictions against the true `q3_test_label`, put it in `q3_score` as `float`/`numpy.float64`. 

To contextualize your model, make a baseline model that always predicts the most frequent class **in the training data**. You can implement this either by yourself or using [`DummyClassifier`](https://scikit-learn.org/stable/modules/generated/sklearn.dummy.DummyClassifier.html). Compute and report the `accuracy_score` of the baseline in `q3b_score`. 

In the written part below, compare the performances of the two models, state whether or not the results are reasonable, and why. Specifically, elaborate why the accuracy score of the baseline model is a lot higher than 1/4. (Hint: you may find [`value_counts`](https://pandas.pydata.org/docs/reference/api/pandas.Series.value_counts.html) helpful.)

```python
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.naive_bayes import MultinomialNB
from sklearn.pipeline import make_pipeline
from sklearn.dummy import DummyClassifier

# Write your code for question 3 here

q3_score = ... # float: test classification accuracy score of the model
q3b_score = ... # float: test classification accuracy score of the baseline


# Leave these lines here to print the result
print("Model test classification accuracy:", q3_score)
print("Baseline test classification accuracy:", q3b_score)
```

<!-- BEGIN QUESTION -->

### Answer 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 4 (6 points, 4 autograded)
In addition to natural language processing, machine learning techniques for preditive modeling can be very effective for analyzing raw image data. In this question, we consider the task of modeling images of handwritten digits. We pose this as a classification task: We want to predict the digit (0, 1, 2, ..., 9) as a function of an 8 by 8 pixel grayscale image of a digit.

To begin, we import and preview the data below. `data` has 64 features, each of which is an integer between 0 and 15 corresponding to a particular pixel intensity value for one of the 64 pixels in the 8 by 8 image. `target` contains the actual digits corresponding to each particular image.

```python
# Run but do not modify this code

import matplotlib.pyplot as plt

# Importing data
from sklearn.datasets import load_digits
digits = load_digits()
q4_data = digits["data"]
q4_target = digits["target"]

# Previewing data
print("The first image corresponds to digit: ", q4_target[0], "\n")
print("First image with pixels flattened out: \n", q4_data[0], "\n")
print("First image rearranged to be 8 x 8: \n", q4_data[0].reshape((8,8)), "\n")
print("Visualizing first several images:")

# Visualizing digits
fig, axes = plt.subplots(3, 10, figsize=(10, 3),
                         subplot_kw={'xticks':[], 'yticks':[]},
                         gridspec_kw=dict(hspace=0.1, wspace=0.1))
for i, ax in enumerate(axes.flat):
    ax.imshow(digits.images[i], cmap='binary')
```

We first split the data into training and test sets as usual.

```python
# Run but do not modify this code

q4_train_data, q4_test_data, q4_train_target, q4_test_target = train_test_split(
    q4_data, q4_target, test_size=0.3, random_state=2)
```

Now, use a k-nearest neighbors classifier for your model that predicts the underlying digit given an image. We import the relevant `KNeighborsClassifier` for you below. Recall that an important hyperparameter of a k-nearest neighbor model is k, or the `n_neighbors` parameter to the `KNeighborsClassifier()` constructor. 

Use 5-fold cross validation to search for the best value of `n_neighbors` in validation on the training set from among possible values between 2 and 20. You can either do so using `cross_validate` to check each possible setting directly, or you can use a `GridSearchCV` to perform the search automatically during fitting. If you search manually, take the average of the five scores for 5-fold cross validation, for each value. If you use `GridSearchCV`, note that the `GridSearchCV` object can itself fit and predict, and that you can get the optimal setting after fitting with `.best_params_`. Report the optimal value of `n_neighbors` in `q4_n` as `int`/`numpy.int64`.

When you have fit your model with the optimal `n_neighbors` value on the training data, make predictions for the test data. Compute and report the `accuracy_score` of your model in testing and report it in `q4_score` as `float`/`numpy.float64`.

Next, fit a Naive Bayes model (using `MultinomialNB` and the default setting) on the training data and predict the class of the test data.  Compute its `accuracy_score` and put it in `q4b_score`. In the written part, compare the two `accuracy_score`, discuss whether or not they are reasonable, and why.

```python
from sklearn.model_selection import cross_validate
from sklearn.model_selection import GridSearchCV
from sklearn.neighbors import KNeighborsClassifier

# Write your code for question 4 here

q4_n = ... # int: n_neighbors
q4_score = ... # float: test accuracy score of the k nearest neighbor model
q4b_score = ... # float: test accuracy score of the naive bayes model


# Leave these lines here to print the result
print("Optimal n_neighors:", q4_n)
print("Accuracy on testing data:", q4_score)
print("Accuracy of baseline:", q4b_score)
```

<!-- BEGIN QUESTION -->

### Answer 4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## AI Disclosure

Use the [Artificial Intelligence Disclosure (AID) Framework](https://crln.acrl.org/index.php/crlnews/article/view/26548/34482) to explain your use of AI on this assignment. Other headings you can use include:
- *Debugging*: Using AI to help you fix your code so that it works. You should state how you used it for this purpose.

Here are some examples:

*Artificial Intelligence Tools*: ChatGPT v5 via chatgpt.com. *Conceptualization*: I gave chatgpt.com the election data set and asked it for ideas on interesting statistics I could get from the data. *Methodology*: I asked it for help on how to write the code to get the statistic I chose, but I wrote the code myself. *Writing—Review & Editing*: I wrote out my explanation for what the statistic meant, then gave that text and the rubric to chatgpt and asked it to give me feedback on how to update the explanation to conform to the rubric.

*Artificial Intelligence Tools*: ChatGPT v4o via  DukeGPT. *Information Collection*: DukeGPT was used to find the function needed to get the index value of the maximum value of a Series and the syntax needed to filter rows in a pandas dataframe using multiple columns. *Debugging*: DukeGPT was used to help me find a bug in my code for Q1 where I copied in the code and error, stated what the code should do, and asked for help.

**Solution:** Answer is here

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **required** to go to the Kernel option and choose **"Restart & Run All"**  before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/ "‌"). **Submissions not prepared correctly as above will lose points.**

Tests for check_bad_variable_names (points: [0.5])
```
>>> from auto_checks import get_default_type_definitions
>>> types_to_check = {'int': int, 'all': all, 'min': min, 'max': max, 'any': any, 'isinstance': isinstance, 'float': float, 'complex': complex, 'str': str, 'bytes': bytes, 'bytearray': bytearray, 'list': list, 'tuple': tuple, 'range': range, 'dict': dict, 'set': set, 'frozenset': frozenset, 'bool': bool, 'NoneType': type(None), 'Ellipsis': type(Ellipsis), 'NotImplemented': type(NotImplemented), 'slice': slice, 'memoryview': memoryview, 'classmethod': classmethod, 'staticmethod': staticmethod, 'property': property, 'object': object}
>>> def are_types_overwritten(types_to_check):
...     """
...     Joseph Scarpa, 02-24-2024
...     
...     Checks if built-in types are overwritten by custom definitions.
...     - types_to_check: A dictionary mapping from type names (as strings) to the custom types defined in the student notebook.
...     
...     Returns:
...     - Boolean indicating whether the built-in types are overridden or not.
...     """
...     built_in_types = get_default_type_definitions()
...     for type_name, type_val in types_to_check.items():
...         if built_in_types.get(type_name) is not type_val:
...             msg = f'NOT YET: It looks like you have created a variable called {type_name}. '
...             msg += 'This is a reserved keyword in Python. Use a different variable name.'
...             print(msg)
...             return False
...     return True
>>> are_types_overwritten(types_to_check)
True
```


Tests for q1_1-code (points: [0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.5, 0.5, 0.5, 0.5])
```
>>> isinstance(q1_mse, float) or isinstance(q1_mse, np.float64)
True
```
```
>>> bool(q1_mse > 0)
True
```
```
>>> isinstance(q1_r2, float) or isinstance(q1_r2, np.float64)
True
```
```
>>> bool(q1_r2 <= 1)
True
```
```
>>> isinstance(q1b_mse, float) or isinstance(q1b_mse, np.float64)
True
```
```
>>> bool(q1b_mse > 0)
True
```
```
>>> isinstance(q1b_r2, float) or isinstance(q1b_r2, np.float64)
True
```
```
>>> bool(q1b_r2 <= 1)
True
```
```
>>> from auto_checks import check_q1_mse
>>> check_q1_mse(q1_mse)
'PASS'
```
```
>>> from auto_checks import check_q1_r2
>>> check_q1_r2(q1_r2)
'PASS'
```
```
>>> from auto_checks import check_q1b_mse
>>> check_q1b_mse(q1b_mse)
'PASS'
```
```
>>> from auto_checks import check_q1b_r2
>>> check_q1b_r2(q1b_r2)
'PASS'
```


Tests for q1_2-code (points: [0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.5, 0.5, 0.5, 0.5])
```
>>> isinstance(q12_1_mse, float) or isinstance(q12_1_mse, np.float64)
True
```
```
>>> bool(q12_1_mse > 0)
True
```
```
>>> isinstance(q12_1_r2, float) or isinstance(q12_1_r2, np.float64)
True
```
```
>>> bool(q12_1_r2 <= 1)
True
```
```
>>> isinstance(q12_2_mse, float) or isinstance(q12_2_mse, np.float64)
True
```
```
>>> bool(q12_2_mse > 0)
True
```
```
>>> isinstance(q12_2_r2, float) or isinstance(q12_2_r2, np.float64)
True
```
```
>>> bool(q12_2_r2 <= 1)
True
```
```
>>> from auto_checks import check_q12_1_mse
>>> check_q12_1_mse(q12_1_mse)
'PASS'
```
```
>>> from auto_checks import check_q12_1_r2
>>> check_q12_1_r2(q12_1_r2)
'PASS'
```
```
>>> from auto_checks import check_q12_2_mse
>>> check_q12_2_mse(q12_2_mse)
'PASS'
```
```
>>> from auto_checks import check_q12_2_r2
>>> check_q12_2_r2(q12_2_r2)
'PASS'
```


Tests for q2-code (points: [0.25, 0.25, 0.5, 0.5, 0.5, 1])
```
>>> isinstance(q2_score, float) or isinstance(q2_score, np.float64)
True
```
```
>>> bool(0 <= q2_score <= 1)
True
```
```
>>> isinstance(q2b_score, float) or isinstance(q2b_score, np.float64)
True
```
```
>>> bool(0 <= q2b_score <= 1)
True
```
```
>>> from auto_checks import check_q2_score
>>> check_q2_score(q2_score)
'PASS'
```
```
>>> from auto_checks import check_q2b_score
>>> check_q2b_score(q2b_score)
'PASS'
```


Tests for q3-code (points: [0.5, 0.5, 0.5, 0.5, 1, 1])
```
>>> isinstance(q3_score, float) or isinstance(q3_score, np.float64)
True
```
```
>>> bool(0 <= q3_score <= 1)
True
```
```
>>> isinstance(q3b_score, float) or isinstance(q3b_score, np.float64)
True
```
```
>>> bool(0 <= q3b_score <= 1)
True
```
```
>>> from auto_checks import check_q3_score
>>> check_q3_score(q3_score)
'PASS'
```
```
>>> from auto_checks import check_q3b_score
>>> check_q3b_score(q3b_score)
'PASS'
```


Tests for q4-code (points: [0.25, 0.25, 0.25, 0.25, 0.5, 0.5, 0.5, 0.5, 1])
```
>>> isinstance(q4_n, int) or isinstance(q4_n, np.int64)
True
```
```
>>> bool(2 <= q4_n <= 20)
True
```
```
>>> isinstance(q4_score, float) or isinstance(q4_score, np.float64)
True
```
```
>>> bool(0 <= q4_score <= 1)
True
```
```
>>> isinstance(q4b_score, float) or isinstance(q4b_score, np.float64)
True
```
```
>>> bool(0 <= q4b_score <= 1)
True
```
```
>>> from auto_checks import check_q4_n
>>> check_q4_n(q4_n)
'PASS'
```
```
>>> from auto_checks import check_q4_score
>>> check_q4_score(q4_score)
'PASS'
```
```
>>> from auto_checks import check_q4b_score
>>> check_q4b_score(q4b_score)
'PASS'
```
