# Worked Example 05
In this exercise, you will demonstrate your learning by solving probability problems. Many of these will be math problems. You can nicely format your answers in Markdown cells using LaTeX-like commands. 

Some particular LaTeX commands you might find useful include:
- To display an inline math expression, enclose the content in single dollar signs, e.g., $x^2+4=8$ (Double click to display).
- To display a single-line centered math expression, enclose the content in double dollar signs like $$P(X > 5) = \frac{1}{6}.$$
- To display a multi-line math expression, use `\begin{aligned}` and `\end{aligned}`, use `\\` to change lines, put `&` wherever you want the lines to align, and enclose all this in single dollar signs. See the example below. Note that this is not normally what people would use in a pure LaTeX environment (it is an overkill), but doing this makes sure your solutions actually render on Gradescope for easy grading.
- Another way to display multi-line math expressions is by enclosing the equation inside double dollar signs. Use `\\` to change lines, `&` is not required here. See example below.
- To write a fraction, use `\frac{numerator}{denominator}` like $\frac{3}{4}$.
- To write set intersection and union, use `\cap` and `\cup` respectively, like $A \cap B$ and $A \cup B$.
- To write a bar over something (like a set complement), use `\overline{}` like in $\overline{E}$.
- To write a subscript or superscript, use `_{}` or `^{}` as in $X_{n}$ or $X^{2}$.
- To write a sum or product, use `\sum` or `\prod` as in $\sum_{i=1}^{n} X_i$ or $\prod_{i=1}^{n} X_i$.
- To write text inside math expressions, use `\text{}` as in $P(D \text{ is even})$.
- You can double click on this cell to get an example of the syntax of the commands mentioned above.

If you are new to LaTeX, see the [overleaf documentation for mathematical expressions](https://www.overleaf.com/learn/latex/mathematical_expressions). Note that not all LaTeX commands work in Markdown cells, as they are technically not actual LaTeX environments. Hence double check in Gradescope after submission.

## Be careful while using LaTeX

- Some of common LaTeX commands like `\begin{equation}`, `\begin{align}` do not render in Gradescope. This is necessary for your homework assignments.
- Even if you are using `\begin{aligned}` and `\end{aligned}`, write everything inside single dollar signs.
- Gradescope mostly renders all LaTeX/Math equations written inside single dollar signs. For example: $\sum_{k=1}^n x_n^2$.

## Show your work!
Show your work and/or briefly explain your answers. To get full credit you must first translate the problem into a probability formula. In general you will not receive full credit for numeric answers with no accompanying work or justification (math, code, explanation), and we can't give you partial credit if you don't show your work.

An example of the level of explanation we are expecting is if we asked you to calculate the probability that a die roll is even or > 3 (like in the videos), we would first expect a translation to probabilistic terms and then transformations until you find the number like so:

Let $D$ be the result of the die roll.

$$ 
P(D \text{ is even} \cup D > 3) = P(D \text{ is even}) + P(D > 3) - P(D \text{ is even} \cap D > 3) \\
= \frac{1}{2} + \frac{1}{2} - \frac{1}{3} = \frac{2}{3}.
 $$


### Autograder Instructions

Run the following cell to install the `Otter` autograder package the first time it is run. Afterwards, running it again will print out its version number if `Otter` is correctly installed.

```python
try:
    import otter
    grader = otter.Notebook("worked-example05.ipynb")
    print(otter.__version__)
except Exception as e:
    %pip install otter-grader
```

Output:
```
6.1.6
```

Unlike Homeworks, the autograders for worked examples are embedded with the questions so that you can run the tests yourself. To grade your own work, simply run the cell starting with `grader.check` immediately after each question. The cell calls Otter to run the tests for all subquestions and generate a report of what test(s) you pass/do not pass in the same style as on Gradescope.

### Question 1
Suppose we roll a single fair six-sided die (meaning we get a number from 1 to 6 drawn uniformly at random) and toss a single fair coin (meaning we get heads or tails each with 50% probability). For each of the following question, what is the probability of the event? Show your work in Markdown.

For the autograder, in the code cell, put question 1's answer in the variable `q1_1`, question 2's answer in `q1_2`, and so on.

1. An odd die roll and a coin toss of heads.
2. A die roll greater than or equal to 4 and anything for the coin.
3. A coin toss of tails and anything for the die.
4. A die roll greater than or equal to 3 or a coin toss of tails (possibly both).

### Answer 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Save results in variables
q1_1 = ...
q1_2 = ...
q1_3 = ...
q1_4 = ...
```

```python
grader.check("q1")
```

### Question 2
Assume 500 students enrolled in both Calculus and Physics. Of these students,
- 82 got an A in Calculus, 
- 73 got an A in Physics, and 
- 42 got an A in both courses.

For each of the following, compute the probability that a student chosen uniformly at random:
1. Got an A in both courses.
2. Got less than an A in at least one of the two courses.
3. Got an A in Calculus but not in Physics.
4. Got an A in at least one of the two courses.

For the autograder, in the code cell, put question 1's answer in the variable `q2_1`, question 2's answer in `q2_2`, and so on.

### Answer 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Save results in variables
q2_1 = ...
q2_2 = ...
q2_3 = ...
q2_4 = ...
```

```python
grader.check("q2")
```

### Question 3
Consider a certain disease. Suppose that the probability $P(D)$ that a random person has the disease is 1% or 0.01.


1. For a random person, suppose the probability $P(T)$ that they get tested for the disease is 20% or 0.2. Also, suppose the probability $P(D \cap T)$ that they both have the disease and get tested is 0.5% or 0.005. What is the probability $P(D \mid T)$ that they have the disease given that they get tested?


2. Suppose the test is not always accurate. In particular, the probability $P(pos \mid D)$ of a positive test result given someone has the disease is 95% or 0.95. The probability $P(pos \mid \overline{D})$ of a positive test result given someone does **not** have the disease is 10% or 0.1. Suppose a random person takes the test. What is the probability $P(pos)$ that they test positive?


3. What is the probability $P(D \mid pos)$ that a random person has the disease given that they test positive?

For the autograder, in the code cell, put answers in the variables `q3_1`, `q3_2`, and so on.

### Answer 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python

q3_1 = 0.005 / 0.2  # P(D and T) / P(T)
q3_2 = 0.95 * 0.01 + 0.1 * 0.99  # P(pos|D)*P(D) + P(pos|~D)*P(~D)
q3_3 = (0.95 * 0.01) / (0.95 * 0.01 + 0.1 * 0.99)  # Bayes rule: P(D|pos)

```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Question 4
Suppose a student is taking a 50-question multiple-choice exam with 4 answer choices per question. For each question, the probability that the student knows the correct answer is 80%. The student correctly answers these questions, and guesses uniformly at random on the other questions.


1. For a single quesiton, what is the student's expected score if correct answers get 2 points and incorrect answers get 0 points?


2. What is the student's expected overall score on all 50 questions if correct answers get 2 points and incorrect answers are penalized with 0.5 negative points?

For the autograder, in the code cell, put answers in the variables `q4_1` and `q4_2`.

### Answer 4

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Save results in variables
q4_1 = ...
q4_2 = ...
```

```python
grader.check("q4")
```

### Question 5
To simplify this problem, assume that a year has twelve equal duration months, rather than the unequal months as in the standard Gregorian calendar. 

1. Assuming that people are equally likely to be born in any month, how many people must be selected independently and uniformly at random before the probability that at least two of the selected people have the same birth month is at least 1/2?


2. Now assume that for each month among May, June, July, and August, the probability of being born in that month is 1/10, whereas for every other month the probability of being born in that month is 3/40. What is the probability that two people selected uniformly at random have the same birth month? Is this higher or lower than in part 1?

For the autograder, in the code cell, put answers in the variables `q5_1` and `q5_2`.

### Answer 5

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Save results in variables
q5_1 = ...
q5_2 = ...
```

```python
grader.check("q5")
```

### Question 6

1. Suppose the average/mean home price in the United States is 300,000 dollars. If we select a home uniformly at random, what is an upper bound on the probability that its price is more than 1 million dollars? That is, provide and justify a statement of the form "The probability that a home chosen uniformly at random has price greater than $X$ is at most $P$." Use Markov's inequality to provide the bound.


2. Suppose the standard deviation in home prices is 100,000 dollars. Use Chebyshev's inequality to provide a better upper bound on the probability that a home selected unformly at random has a price of more than 1 million dollars.


3. Suppose we have an iid sample of $n$ home prices and the empirical mean price $\overline{X}_n$ is greater than 301,000 dollars. Assuming that we know the standard deviation in home prices is 100,000 dollars and assuming $n$ is large enough that the central limit theorem applies, how many samples $n$ would we need before we could conclude that the probability of observing $\overline{X}_n > 301,000$ is at most 5%, assuming the true mean of the underlying distribution is 300,000 dollars?

For the autograder, in the code cell, put answers in the variables `q6_1`, `q6_2`, and `q6_3`.

### Answer 6

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Save results in variables
q6_1 = ...
q6_2 = ...
q6_3 = ...
q6_1 = 0.3  # By Markov's inequality: P(X > 1,000,000) ≤ 300,000/1,000,000 = 0.3
q6_2 = 0.04  # By Chebyshev's: P(|X-300,000| ≥ 700,000) ≤ (100,000/700,000)^2 ≈ 0.0204 ⇒ for ≥+700k, bound is double-sided so use 2.04% *2 = 4%
q6_3 = 400  # For CLT, z = (301000-300000)/(100000/sqrt(n)), need P(Z>1/std error)≤0.05 ⇒ Z≥1.645, so n ≥ (1000/(100000/√n))=1.645 ⇒ n ≥ (1000/(1.645*100000))^2 ≈ 400

```

```python
grader.check("q6")
```

Output:
```
q6 results:
    q6 - 1 result:
        ✅ Test case passed

    q6 - 2 result:
        ❌ Test case failed
        Trying:
            bool(round(q6_2, 5) == 0.02041)
        Expecting:
            True
        **********************************************************************
        Line 1, in q6 1
        Failed example:
            bool(round(q6_2, 5) == 0.02041)
        Expected:
            True
        Got:
            False

    q6 - 3 result:
        ❌ Test case failed
        Trying:
            q6_3 in range(36100, 40805)
        Expecting:
            True
        **********************************************************************
        Line 1, in q6 2
        Failed example:
            q6_3 in range(36100, 40805)
        Expected:
            True
        Got:
            False
```

### Question 7
There were 151 "pokemon" in the original generation of the popular franchise. Suppose you start purchasing pokemon cards. Each card you purchase has a pokemon chosen independently and uniformly at random from the 151 possible with replacement (that is, you can get the same card multiple times). Use Monte Carlo Simulation to estimate answers to the following questions. You can use Python's `random` library or the `numpy.random` library.

1. Suppose you have a particular favorite pokemon. Suppose you purchase 150 cards total. What is the probability that you will get your favorite pokemon's card? Conduct enough simulations that you are confident in your answer +/- 0.01 (or 1%).


2. Suppose you purchase 300 cards total. What is the expected number of *unique* pokemon cards that you will have? Conduct enough simulations that you are confident in your answer +/- 1 card.


3. Suppose you buy 600 cards total. What is the probability that you will have all 151 unique cards? Conduct enough simulations that you are confident in your answer +/- 0.01 (or 1%).

For the autograder, in the code cell, put answers in the variables `q7_1`, `q7_2`, and `q7_3`.

<!-- BEGIN QUESTION -->



```python
# Put your code to answer question 7 here, and save results in variables below
# Feel free to add extra code cells as needed
```

```python

```

```python

```

```python

```

<!-- END QUESTION -->



```python
# Save results in variables
q7_1 = ...
q7_2 = ...
q7_3 = ...
```

```python
grader.check("q7")
```

### Autograder Instructions

Unlike Homeworks, the autograders for worked examples are embedded with the questions so that you can run the tests yourself. To grade your own work, simply run the cell starting with `grader.check` immediately after each question. The cell calls Otter to run the tests for all subquestions and generate a report of what test(s) you pass/do not pass in the same style as on Gradescope.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: [2, 2, 2, 2])
```
>>> bool(round(q1_1, 5) == 0.25)
True
```
```
>>> bool(round(q1_2, 5) == 0.5)
True
```
```
>>> bool(round(q1_3, 5) == 0.5)
True
```
```
>>> bool(round(q1_4, 2) == 0.83)
True
```


Tests for q2 (points: [2, 2, 2, 2])
```
>>> bool(round(q2_1, 5) == 0.084)
True
```
```
>>> bool(round(q2_2, 5) == 0.916)
True
```
```
>>> bool(round(q2_3, 5) == 0.08)
True
```
```
>>> bool(round(q2_4, 5) == 0.226)
True
```


Tests for q3 (points: [2, 2, 2])
```
>>> bool(round(q3_1, 5) == 0.025)
True
```
```
>>> bool(round(q3_2, 5) == 0.1085)
True
```
```
>>> bool(round(q3_3, 5) == 0.08756)
True
```


Tests for q4 (points: [2, 2])
```
>>> bool(round(q4_1, 5) == 1.7)
True
```
```
>>> bool(round(q4_2, 5) == 81.25)
True
```


Tests for q5 (points: [2, 2])
```
>>> bool(q5_1 == 5)
True
```
```
>>> bool(round(q5_2, 5) == 0.085)
True
```


Tests for q6 (points: [2, 2, 2])
```
>>> bool(round(q6_1, 5) == 0.3)
True
```
```
>>> bool(round(q6_2, 5) == 0.02041)
True
```
```
>>> q6_3 in range(36100, 40805)
True
```


Tests for q7 (points: [2, 2, 2])
```
>>> bool(0.62 <= q7_1 <= 0.64)
True
```
```
>>> bool(129 <= q7_2 <= 131)
True
```
```
>>> bool(0.04 <= q7_3 <= 0.06)
True
```
