# Worked Example 09

In this exercise, you will practice working with a database using SQL. All questions can be done entirely in SQL without manipulating the data after querying it from the database. If you cannot figure out the SQL and find it easier to manipulate the data in pandas, you may do so, but you must use SQL to extract the data from the database.

### Autograder Instructions

Run the following cell to install the `Otter` autograder package the first time it is run. Afterwards, running it again will print out its version number if `Otter` is correctly installed.

```python
try:
    import otter
    grader = otter.Notebook("worked-example09.ipynb")
    print(otter.__version__)
    if (otter.__version__ < '5.2.2'): #update for latest otter version
        %pip install -U --user otter-grader
except Exception as e:
    %pip install otter-grader
```

## The congress database
A visual schema of the congress database is also available as a pdf along with this practice - we encourage you to start by taking a look to get familiar. This database contains the history of the members of the United States congress through the 115th congress (the data end by 2019) as well as a good deal of voting data from 2015-2016. The visual schema shows each table in the database as a yellow box with the table name at the top and the column names listed below. The arrows showing which keys/identifiers match between different tables for join operations. At the bottom you can see previews of the tables. One in particular, the `cur_members` table, contains data about members of the 115th congress (note - everywhere you see us refer to "current" members below, we are referring to those in the `cur_members` table). Many, though not all, of these members are still serving as of the 117th congress, which began in January 2021. 

## How to use SQL and report your results

You should use Python sqlite3 and Pandas to run your SQL queries in this assignment. You should include your code (including the Strings containing your SQL queries) in this notebook and show your results for the manual grading. We recommend using multiline strings for your queries for readability; remember that multiline strings in Python are enclosed by triple quotes for example, 

```
my_multiline_string = """Hi 
There!"""
```

We provide an example below.

```python
# Run but do not modify the following code
# to import sqlite3 and pandas, and to connect
# to the congress database
import sqlite3
import pandas as pd
import numpy as np
conn = sqlite3.connect("congress")
```

```python
# Here is an example of how you can write and
# execute queries in a notebook using Pandas. 
# It's the same as the example query written
# above. Feel free to change this cell
# to explore the database.
query = """SELECT id, first_name, last_name
FROM cur_members
WHERE type='sen'
    AND party='Democrat'
LIMIT 3;
"""
pd.read_sql(query, conn)
```

### Question 1
Who are the five youngest members of the `cur_members` table? Show all the information about them from the `cur_members` table.

```python
# Start with a SQL query to extract the information
query1 = """ ... """

# Execute queries using pandas
q1 = ...


# Leave this line here so that the table is shown
q1
```

```python
grader.check("q1")
```

### Question 2
List all past and present female (`gender = 'F'`) members of the Congress who were born in the 1970s. Give their `first_name`, `last_name`, and `birthday`. This will require you to compare dates. In sqlite, the standard date format is `'YYYY-MM-DD'` for year, month, and date (see, for example, the format of the `birthday`s in the database). For advanced date functionalities, see the [date and time documentation](https://sqlite.org/lang_datefunc.html), but for these problems you should be able to compare dates simply using a string in the standard format, for example `>= '1900-01-01'` should be true for birthdays in the year `1900` or later.

```python
# Start with a SQL query to extract the information
query2 = """ ... """

# Execute queries using pandas
q2 = ...


# Leave this line here so that the table is shown
q2
```

```python
grader.check("q2")
```

### Question 3
How many bills (from the `bills` table) have a `short_title`? You can check for a blank `short_title` by checking for the empty string `''`. Put your answer in the variable `q3` as an `int`.

```python
# Start with a SQL query to extract the information
query3 = """ ... """

# Execute queries using pandas & check for the empty string
q3 = ...


# Leave this line here so that the table is shown
q3
```

```python
grader.check("q3")
```

### Question 4
Which five states have the most representatives in the `cur_members` table, and how many representatives do they have? Create a table with two columns `state` and `count`. To break ties, order by the `state` column. Note you'll need to at least rename the `count` column to do this.

```python
# Start with a SQL query to extract the information
query4 = """ ... """

# Execute queries using pandas
q4 = ...


# Leave this line here so that the table is shown
q4
```

```python
grader.check("q4")
```

### Question 5
One of the important votes cast is for electing the Speaker of the House. You can find these votes by looking for `Election of the Speaker` in the `question` column of the `votes` table - there should be two, one in which John Boehner won and another in which Paul Ryan won. 

Consider the 2015 election where the `result` was `Ryan (WI)`. Count how many votes were received by everyone who received votes in that 2015 election won by Paul Ryan. For each, show the `vote` from the `person_votes` table (which is the candidate for whom someone voted) and the count of the number of votes they received. There are 7 such results from `vote`, including `Not Voting`. The values should add up to 435 (the number of voting members in the US House of Representatives).

The answer should be in the variable `q5` as a table with the columns `vote` and `count`. Sort by `vote` (so it's alphabetical) where the vote earliest in the alphabet is at the top.

```python
# Start with a SQL query to extract the information
query5 = """ ... """

# Execute queries using pandas
q5 = ...


# Leave this line here so that the table is shown
q5
```

```python
grader.check("q5")
```

### Question 6
List all North Carolina (`NC`) `Republican` senators (`sen`) past and present by `first_name` and `last_name`, together with the `start_date` of their earliest term as senator and the `end_date` of their latest term (which may or may not be the same as their first depending on if they served multiple terms). Order the results by the `start_date` of their earliest term in descending order (most recent first), break ties using their first name (alphabetical).

The answer should be in the variable `q6` as a table with the columns `first_name`, `last_name`, `party`, `start`, and `end`. Sort by `start` where the most recent start date is at the top.

For example, the first two rows of the result should be the two current North Carolina Senators shown below. Note that although Richard Burr has served multiple terms, we only display one row for him containing his earliest `start_date` and latest `end_date`.
```
first_name|last_name|party|start|end
Thom|Tillis|Republican|2015-01-06|2021-01-03
Richard|Burr|Republican|2005-01-04|2023-01-03
```

Note: This query requires joining two tables and has multiple pieces to it. We suggest you slowly build your query up into what it needs to be rather than try to write the entire thing at once.

```python
# Start with a SQL query to extract the information
query6 = """ ... """

# Execute queries using pandas
q6 = ...


# Leave this line here so that the table is shown
q6
```

```python
grader.check("q6")
```

### Question 7
Find the past and present members of congress who have served in the House (`rep`) and the Senate (`sen`), both representing North Carolina (`NC`). The answer should be in the variable `q7` as a table with the columns `id` from the `persons` table, `first_name`, and `last_name` ordered by `last_name` from A to Z.

Hint: You can use `SELECT COUNT(DISTINCT column)` to find the number of distinct values on the given column. Also this is another case of slowly building your query up into what it needs to be.

```python
# Start with a SQL query to extract the information
query7 = """ ... """

# Execute queries using pandas
q7 = ...


# Leave this line here so that the table is shown
q7
```

```python
grader.check("q7")
```

### Question 8
In the election from question 5, how many members in each state voted for Paul Ryan for this position? The values should add up to the number of votes Paul Ryan received in total from the previous question.

The answer should be in the variable `q8` as a table with the columns `state` and `count` (you'll need to rename it). The table should be ordered first by `count` from greatest to least number of votes and then by `state` from A to Z.

```python
# Start with a SQL query to extract the information
query8 = """ ... """

# Execute queries using pandas
q8 = ...


# Leave this line here so that the table is shown
q8
```

```python
grader.check("q8")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: [1, 1, 2])
```
>>> print(sorted(q1.columns))
['birthday', 'first_name', 'gender', 'id', 'last_name', 'party', 'religion', 'state', 'type']
```
```
>>> len(q1)
5
```
```
>>> print(q1[sorted(q1.columns)].sort_values(by='birthday').to_string(index=False))
  birthday first_name gender      id     last_name      party religion state type
1981-04-12      Tulsi      F G000571       Gabbard   Democrat             HI  rep
1982-05-07       Matt      M G000578         Gaetz Republican             FL  rep
1983-09-12       Trey      M H001074 Hollingsworth Republican             IN  rep
1984-03-03       Mike      M G000579     Gallagher Republican             WI  rep
1984-07-02      Elise      F S001196      Stefanik Republican             NY  rep
```


Tests for q2 (points: [1, 1, 2])
```
>>> print(sorted(q2.columns))
['birthday', 'first_name', 'last_name']
```
```
>>> len(q2)
12
```
```
>>> print(q2[sorted(q2.columns)].sort_values(by='birthday').to_string(index=False))
  birthday first_name       last_name
1970-06-08  Gabrielle        Giffords
1970-07-01       Joni           Ernst
1970-12-03  Stephanie Herseth Sandlin
1971-11-30     Kristi            Noem
1975-10-01      Grace            Meng
1975-12-06        Mia            Love
1976-07-12    Kyrsten          Sinema
1976-07-27     Martha            Roby
1976-08-05  Jenniffer  González-Colón
1976-09-15    Nanette        Barragán
1978-09-16  Stephanie          Murphy
1978-11-03      Jaime Herrera Beutler
```


Tests for q3 (points: [2, 2])
```
>>> isinstance(q3, np.int64) or isinstance(q3, int)
True
```
```
>>> q3
8689
```


Tests for q4 (points: [1, 1, 2])
```
>>> len(q4)
5
```
```
>>> print(sorted(q4.columns))
['count', 'state']
```
```
>>> print(q4[['state', 'count']].sort_values(by=['count', 'state']).to_string(index=False))
state  count
   IL     20
   FL     29
   NY     29
   TX     38
   CA     55
```


Tests for q5 (points: [1, 1, 2])
```
>>> len(q5)
7
```
```
>>> print(sorted(q5.columns))
['count', 'vote']
```
```
>>> print(q5[['vote', 'count']].to_string(index=False))
        vote  count
Colin Powell      1
      Cooper      1
       Lewis      1
  Not Voting      3
      Pelosi    184
   Ryan (WI)    236
Webster (FL)      9
```


Tests for q6 (points: [1, 1, 2])
```
>>> len(q6)
17
```
```
>>> print(sorted(q6.columns))
['end', 'first_name', 'last_name', 'party', 'start']
```
```
>>> print(q6[['first_name', 'last_name', 'party', 'start', 'end']].to_string(index=False))
first_name  last_name      party      start        end
      Thom     Tillis Republican 2015-01-06 2021-01-03
   Richard       Burr Republican 2005-01-04 2023-01-03
 Elizabeth       Dole Republican 2003-01-07 2009-01-03
    Duncan  Faircloth Republican 1993-01-05 1999-01-03
     James   Broyhill Republican 1986-07-14 1986-11-04
      John       East Republican 1981-01-05 1986-06-29
     Jesse      Helms Republican 1973-01-03 2003-01-03
     Jeter  Pritchard Republican 1895-12-02 1903-03-03
      John       Pool Republican 1868-01-01 1873-03-03
    Joseph     Abbott Republican 1868-01-01 1871-03-03
  Montfort     Stokes Republican 1816-01-01 1823-03-03
 Nathaniel      Macon Republican 1815-12-04 1825-03-03
     James     Turner Republican 1805-12-02 1817-03-03
     David      Stone Republican 1801-12-07 1815-03-03
     Jesse   Franklin Republican 1799-12-02 1813-03-03
   Timothy Bloodworth Republican 1795-12-07 1801-03-03
 Alexander     Martin Republican 1793-12-02 1799-03-03
```


Tests for q7 (points: [1, 1, 2])
```
>>> len(q7)
18
```
```
>>> print(sorted(q7.columns))
['first_name', 'id', 'last_name']
```
```
>>> print(q7[['id', 'first_name', 'last_name']].to_string(index=False))
     id first_name  last_name
B000456        Asa      Biggs
B000563    Timothy Bloodworth
B000763       John     Branch
B000966      James   Broyhill
B001135    Richard       Burr
C000524     Thomas   Clingman
E000211     Samuel      Ervin
F000344      Jesse   Franklin
H000679      Clyde       Hoey
L000240      Alton     Lennon
M000034  Nathaniel      Macon
M000096     Willie     Mangum
M000993    Cameron   Morrison
R000144      David       Reid
S000415  Furnifold    Simmons
S000955      David      Stone
U000005    William    Umstead
V000021    Zebulon      Vance
```


Tests for q8 (points: [1, 1, 2])
```
>>> len(q8)
44
```
```
>>> print(sorted(q8.columns))
['count', 'state']
```
```
>>> print(q8[['state', 'count']].to_string(index=False))
state  count
   TX     23
   CA     14
   FL     13
   PA     13
   OH     12
   GA     10
   MI      9
   NC      9
   NY      9
   IL      8
   IN      7
   TN      7
   VA      7
   AL      6
   MO      6
   NJ      6
   SC      6
   LA      5
   OK      5
   AR      4
   AZ      4
   CO      4
   KS      4
   KY      4
   UT      4
   WA      4
   WI      4
   IA      3
   MN      3
   MS      3
   NV      3
   WV      3
   ID      2
   NE      2
   AK      1
   MD      1
   ME      1
   MT      1
   ND      1
   NH      1
   NM      1
   OR      1
   SD      1
   WY      1
```
