# Homework 07
## Completed by: Mohammad Zoraiz

In this exercise, you will practice inferential statistics with confidence intervals, bootstrapping, and hypothesis testing. Problems may involve a combination of math and code. 

Recall that you can use LaTeX to nicely format your math inside Markdown cellsby enclosing equations in single dollar signs (e.g., $x^2+4=8$) for inline math or double dollar signs for centered equations like $$P(X > 5) = \frac{1}{6}.$$ See Worked Examples/HW 05 for further instructions and more examples.

Show your work and/or briefly explain your answers. In general, you will not receive full credit for numeric answers with no accompanying work or justification (math, code, explanation). For numeric answers, we will accept answers that are very slightly off due to rounding, $z$-score of $2$ vs. $1.96$, etc. 

### Scipy Version Check
The autograder uses `scipy` version 1.9.0+, if you do not have this version or higher, you will fail autograder tests. This cell is to ensure you have a high enough version.

```python
try:
    import scipy
    print(scipy.__version__)
    if (scipy.__version__ < '1.10.0'): 
        %pip install --upgrade scipy
except Exception as e:
    %pip install scipy
```

Output:
```
1.13.1
```

```python
# Run this code cell to import relevant libraries
import numpy as np
import pandas as pd
from scipy import stats
```

### Question 1 (4 points, all autograded)

A website is trying to increase registration for first-time visitors, exposing a random subset of these visitors to a new site design. Of $752$ randomly sampled visitors over a month who saw the new design, $64$ registered. Using [`stats.norm.interval`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html), construct a $95\%$ confidence interval for the percentage of visitors who would register for the website under the new design. 

Save your answer in a tuple `q1_1` of two `float` or `numpy.float64` variables such that `q1_1[0]` is the lower/left bound and `q1_1[1]` is the upper/right bound of your confidence interval. We accept either percentages (e.g., $50.0$) or fractions (e.g., $0.5$). [`stats.norm.interval`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html) when used correctly already returns a tuple.

```python
# Code for question 1
q1 = stats.norm.interval(0.95, loc=64/752, scale=np.sqrt((64/752)*(1-64/752)/752))

# Leave these lines here for grading and ease of debugging
print("Q1: 95% confidence interval: ", q1)
```

Output:
```
Q1: 95% confidence interval:  (np.float64(0.06516269200219607), np.float64(0.10505007395525073))
```

### Question 2 (8 points: 4 autograded, 4 manually graded)

A study examined the mean pay for a random sample of men and women entering the workforce as doctors for $43$ different positions. If each gender was equally paid, then we would expect about half of those positions to have men paid more than women and women would be paid more than men in the other half of positions. In the study, men were, on average, paid more in $29$ of the $43$ positions. 

Complete a hypothesis test using `stats.norm.cdf` to examine whether there is significant evidence (at the $0.05$ level) of gender discrimination in pay in these positions. It can be either one-sided or two-sided. Uncomment the corresponding line to indicate which test you are conducting in the variable `q2_setting`. Save your p-value in `q2` as `float` or `numpy.float64`. 

In the written part:
  1. Explicitly write down the null hypothesis and the alternative hypothesis for this test.
  1. Report your p-value and interpret the result.
  1. Explain why using `stats.norm.cdf` (not `stats.t.cdf`) is a good choice for this context.

```python
# Code for question 2
# q2_setting = 1 # Uncomment for one-sided test
# q2_setting = 2 # Uncomment for two-sided test
q2 = 2 * (1 - stats.norm.cdf((29 - 43/2) / np.sqrt(43/4)))
q2_setting = 2  # Two-sided test based on the calculation

# Leave these lines here for grading and ease of debugging
print("Q2: p-value = {} for {}-sided test".format(q2, q2_setting))
```

Output:
```
Q2: p-value = 0.02216790575298888 for 2-sided test
```

<!-- BEGIN QUESTION -->

### Answer 2

The null hypothesis for this test is that there is no gender discrimination in pay among the 43 positions, which means men and women would be paid more than the other in about half of the positions each. The alternative hypothesis is that there is a significant difference in pay, so men are paid more in significantly more than half of the positions.

Based on the calculation, 29 out of 43 positions had men being paid more. Using a two-sided test as shown in the code above, the p-value we calculated is quite small, around 0.0222. This value is below the standard significance level of 0.05. Because the p-value is less than 0.05, we have enough evidence to reject the null hypothesis. This provides significant evidence to suggest that men tend to be paid more than women in these positions and that this result is unlikely to be explained just by chance.

We use stats.norm.cdf (instead of stats.t.cdf) because the central limit theorem tells us that, with 43 observations, the distribution of the number of times men are paid more than women can be reasonably approximated by a normal distribution. Since we are counting the number of positions with men paid more out of a fixed sample size, and this follows a binomial distribution which can be approximated by the normal distribution for large samples, stats.norm.cdf is appropriate here.

Note: The calculated p-value for this test is approximately 0.0222, which is less than the significance level of 0.05. Therefore, we have enough evidence to reject the null hypothesis and conclude that there is significant evidence for a difference (favoring men) in pay in these positions.

<!-- END QUESTION -->

## Movie Ratings Data
In the remainder of this assignment you will work with the movielens dataset of movie ratings that we have seen before. Below we import and preview the data. It consists of 2 tables: `users` has a row for every individual who has rated any movies, `movie-ratings` has a row for every rating of a particular movie by a particular user. This means users with multiple ratings are in the `movie_ratings` multiple times. The data is a random sample of all of the movie ratings made on the movielens service.

```python
users = pd.read_csv("users.csv")
users.head()
```

Output:
```
user_id  age sex  occupation
0        1   24   M  technician
1        2   53   F       other
2        3   23   M      writer
3        4   24   M  technician
4        5   33   F       other
```

```python
movie_ratings = pd.read_csv("movies-all.csv")
movie_ratings.head()
```

Output:
```
user_id  age sex  occupation  movie_id  rating                 movie_title
0        1   24   M  technician        61       4  Three Colors: White (1994)
1       13   47   M    educator        61       4  Three Colors: White (1994)
2       18   35   F       other        61       4  Three Colors: White (1994)
3       58   27   M  programmer        61       5  Three Colors: White (1994)
4       59   49   M    educator        61       4  Three Colors: White (1994)
```

### Question 3 (12 points: 8 autograded, 4 manually graded)

1. Compute a $95\%$ confidence interval for the mean `age` of users using `stats.norm.interval`. Save your answer in a tuple `q3_1` similar to that in Q1.
2. Compute a $95\%$ confidence interval for the mean `age` of users **who have rated the movie `Casablanca (1942)`** using the normal distribution. Save your answer in a tuple `q3_2`.

In the written part, answer the following question.

3. `Casablanca (1942)` is an old movie, one might suspect that it has been rated by older individuals on average than the entire dataset. Just looking at the confidence intervals you computed in steps 1 and 2, can you conclude that there is significant evidence for this belief? Why or why not? 

```python
# Code for question 3
# Code for question 3

q3_1 = stats.norm.interval(0.95, loc=users['age'].mean(), scale=users['age'].std(ddof=1)/np.sqrt(len(users)))
q3_2 = stats.norm.interval(
    0.95,
    loc=users.merge(movie_ratings[movie_ratings['movie_title']=='Casablanca (1942)'], on='user_id')['age_x'].mean(),
    scale=users.merge(movie_ratings[movie_ratings['movie_title']=='Casablanca (1942)'], on='user_id')['age_x'].std(ddof=1) /
          np.sqrt(users.merge(movie_ratings[movie_ratings['movie_title']=='Casablanca (1942)'], on='user_id')['age_x'].count())
)

# Leave these lines here for grading and ease of debugging
print("3.1: 95% confidence interval: ", q3_1)
print("3.2: 95% confidence interval: ", q3_2)
```

Output:
```
3.1: 95% confidence interval:  (np.float64(33.27375766393051), np.float64(34.830165984001624))
3.2: 95% confidence interval:  (np.float64(34.460497318851566), np.float64(37.33374136427601))
```

<!-- BEGIN QUESTION -->

## Answer 3

Based on the confidence intervals, the mean age of all users is between about 33.27 and 34.83 years, while for users who rated Casablanca (1942), the mean age is between about 34.46 and 37.33 years with 95% confidence. Because these intervals overlap slightly, we cannot confidently say that people who rated Casablanca are definitely older—there is some evidence that their average age might be higher, but the overlap means we do not have strong enough evidence to draw a firm conclusion just from these confidence intervals. There could still be a difference, but the data alone here is not enough to prove it.

<!-- END QUESTION -->

### Question 4 (6 points, all autograded)
Only $18$ users have rated the movie `Lost in Space (1998)`.
1. Use bootstrapping with $10000$ bootstrap resamples to compute a $95\%$ confidence interval for the mean `age` of users who have rated `Lost in Space (1998)`. Save your answer in a tuple `q4_1`.
2. One of the advantages of bootstrapping is that we can easily compute confidence intervals for arbitrary measurements of distributions. Use bootstrapping with $10000$ bootstrap resamples to compute a $95\%$ confidence interval for the **median** `rating` of `Lost in Space (1998)`. Note that `numpy` provides a vectorized function for [calculating the median](https://numpy.org/doc/stable/reference/generated/numpy.median.html) as well as the mean. Save your answer in a tuple `q4_2`.

```python
# # Code for question 4

lost_users = users[users['user_id'].isin(
    movie_ratings[movie_ratings['movie_title']=='Lost in Space (1998)']['user_id']
)]['age'].to_numpy()
lost_ratings = movie_ratings[movie_ratings['movie_title']=='Lost in Space (1998)']['rating'].to_numpy()

rng = np.random.default_rng(42)

q4_1 = tuple(np.percentile(rng.choice(lost_users, size=(10000, len(lost_users)), replace=True).mean(axis=1), [2.5, 97.5]))
q4_2 = tuple(np.percentile(np.median(rng.choice(lost_ratings, size=(10000, len(lost_ratings)), replace=True), axis=1), [2.5, 97.5]))



# Leave these lines here for grading and ease of debugging
print("4.1: 95% confidence interval: ", q4_1)
print("4.2: 95% confidence interval: ", q4_2)
```

Output:
```
4.1: 95% confidence interval:  (np.float64(25.944444444444443), np.float64(36.77777777777778))
4.2: 95% confidence interval:  (np.float64(2.5), np.float64(4.0))
```

### Question 5 (15 points: 6 autograded, 9 manually graded)

1. (5pts) Consider the null hypothesis that the mean rating of `Muppet Treasure Island (1996)` is the same for `sex='F'` and `sex='M'` users. The alternative hypothesis is that the mean ratings are not equal. Conduct a two-sided t test using [`stats.ttest_ind`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind.html) to evaluate this using the sample ratings data. Report your p-value in `q5_1` as a `float` or `numpy.float64`. Interpret your p-value at a significance level of $0.05$ in the **Answer 5** cell. **Moreover, explain the underlying assumptions in your t-test in the cell and whether how you ran the function is reasonable given those assumptions including empirical values from the data as needed.** You are welcome to consult the documentation of [`stats.ttest_ind`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind.html) to help answer this question.

The `Star Wars (1977)` film is quite popular, with a median rating of $5/5$. However, of those that left a rating, male users gave it a slightly higher mean rating of about $4.4$ whereas female users gave the same movie a mean rating of about $4.2$.

2. (4pts) Consider the null hypothesis that $51\%$ of men would rate `Star Wars (1977)` a $5$. The alternative hypothesis is then that the fraction of men who would rate `Star Wars (1977)` a $5$ is not $51\%$. Conduct a two-sided hypothesis test using [`stats.t.cdf`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.t.html) to evaluate this in light of the sample ratings data of male users who rated `Star Wars (1977)`. Report your p-value in `q5_2` as a `float` or `numpy.float64`. Interpret your p-value at a significance level of $0.05$ in the **Answer 5** cell.


3. (4pts) Consider the null hypothesis that women and men were equally likely to rate `Star Wars (1977)` a $5$. The alternative hypothesis is that these proportions are not equal. Conduct a two-sided t test using [`stats.ttest_ind_from_stats`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_ind_from_stats.html) to evaluate this in light of the sample data of female and male users who rated `Star Wars (1977)`. Report your p-value in `q5_3` as a `float` or `numpy.float64`. Interpret your p-value at a significance level of $0.05$ in the **Answer 5** cell. 


4. (2pts) In the sample ratings data, $51\%$ of women rated `Star Wars (1977)` a $5$. (You are welcome to write code to validate this claim, but you can just accept it as fact.) Therefore, if $51\%$ of men and $51\%$ of women would rate `Star Wars (1977)` a $5$, both null hypotheses in 5.2 and 5.3 would be true. If this is the case, why do you observe a different p-value in the two parts, despite the hypotheses under consideration being ostensibly similar? Briefly explain why you observe this difference in the **Answer 5** cell. 

```python
# Code for question 5
q5_1 = stats.ttest_ind(movie_ratings[(movie_ratings['movie_title']=='Muppet Treasure Island (1996)') & (movie_ratings['sex']=='F')]['rating'], movie_ratings[(movie_ratings['movie_title']=='Muppet Treasure Island (1996)') & (movie_ratings['sex']=='M')]['rating'], equal_var=False).pvalue
q5_2 = 2 * (1 - stats.t.cdf((movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].eq(5).mean() - 0.51) / (np.sqrt(0.51 * 0.49 / movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].shape[0])), df=movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].shape[0]-1))
q5_3 = stats.ttest_ind_from_stats(mean1=movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='F')]['rating'].eq(5).mean(), std1=np.sqrt(movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='F')]['rating'].eq(5).mean() * (1-movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='F')]['rating'].eq(5).mean())), nobs1=movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='F')]['rating'].shape[0], mean2=movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].eq(5).mean(), std2=np.sqrt(movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].eq(5).mean() * (1-movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].eq(5).mean())), nobs2=movie_ratings[(movie_ratings['movie_title']=='Star Wars (1977)') & (movie_ratings['sex']=='M')]['rating'].shape[0], equal_var=False).pvalue

# Leave these lines here for grading and ease of debugging
print("Q5.1: p-value: ", q5_1)
print("Q5.2: p-value: ", q5_2)
print("Q5.3: p-value: ", q5_3)
```

Output:
```
Q5.1: p-value:  0.07920295536927348
Q5.2: p-value:  0.008010296971218134
Q5.3: p-value:  0.17469334777160547
```

### Answer 5


Q5.1: For question 5.1, the p-value is about 0.08. Since this is greater than the significance level of 0.05, we do not have enough evidence to reject the null hypothesis. This means that we do not find a significant difference in the mean ratings of "Muppet Treasure Island (1996)" between female and male users given our sample data. The assumptions behind the t-test include that the samples are independent, the ratings are approximately normally distributed in each group, and the variances are not necessarily equal (since we used equal_var=False), which is reasonable for user ratings data since group sizes might be different and variances could be unequal. If the data appear to be skewed or heavily non-normal, the result should be interpreted with caution, but with reasonably large group sizes, the test is fairly robust.

Q5.2: In question 5.2, the p-value is about 0.008. This is much less than the significance level of 0.05, so we reject the null hypothesis that 51% of men would rate "Star Wars (1977)" a 5. Our sample suggests that the true proportion of male users giving a 5 may be different from 51%. The calculation uses the sample proportion, the hypothesized proportion, and estimates the standard error using the null hypothesis value.

Q5.3: For question 5.3, the p-value is about 0.17. This is larger than 0.05, so we do not have enough evidence to reject the null hypothesis that men and women were equally likely to rate "Star Wars (1977)" a 5. This test uses sample means and estimated standard deviations for proportions of 5-star ratings for both men and women, and calculates the t-statistic based on these summary statistics.

For part 4, the difference in p-values between questions 5.2 and 5.3 comes from the fact that question 5.2 tests whether the proportion for men is exactly 51%, based only on the sample of men, while question 5.3 compares the proportions for men and women directly, using data from both groups. Even if both true proportions are 51%, the amount of evidence provided by each test will depend on the sample sizes and observed differences, and the tests are sensitive to different things. In practice, the sample statistics may not be exactly equal to what the null hypotheses specify, so the results can differ.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# AI Disclosure

Use the **Artificial Intelligence Disclosure (AID) Framework** to explain your use of AI on this assignment. Other headings you can use include:

- *Debugging:* Using AI to help you fix your code so that it works. You should state how you used it for this purpose.

Here are some examples:

*Artificial Intelligence Tools:* ChatGPT v5 via chatgpt.com. *Conceptualization:* I gave chatgpt.com the election data set and asked it for ideas on interesting statistics I could get from the data. *Methodology:* I asked it for help on how to write the code to get the statistic I chose, but I wrote the code myself. *Writing — Review & Editing:* I wrote out my explanation for what the statistic meant, then gave that text and the rubric to chatgpt and asked it to give me feedback on how to update the explanation to conform to the rubric.

*Artificial Intelligence Tools:* ChatGPT v4o via DukeGPT. *Information Collection:* DukeGPT was used to find the function needed to get the index value of the maximum value of a Series and the syntax needed to filter rows in a pandas dataframe using multiple columns. *Debugging:* DukeGPT was used to help me find a bug in my code for Q1 where I copied in the code and error, stated what the code should do, and asked for help.

Artificial Intelligence Tools: ChatGPT (various versions) and Claude Opus 4.5

Conceptualization: I used ChatGPT to clarify the intent behind each question and to confirm my understanding of how to approach the problems.

Methodology: I consulted ChatGPT to confirm my use of pandas functions and to ensure I was meeting assignment requirements (such as not using for-loops), then wrote the final code myself.

Debugging: My primary use of ChatGPT was for debugging, especially for Question 4. I relied on it to help resolve issues with the way I was bootstrapping my statistics, including clarifying how to properly perform the resampling, structuring the code, and tracking down related errors. I described the errors I encountered or what my code was supposed to do, and ChatGPT provided suggestions for corrections not only for the bootstrapping function but also for associated pandas and numpy mistakes.

Writing – Review & Editing: For written answers, I asked ChatGPT for suggestions to make my responses—especially my answer to Question 5 in the pivot table section—more concise and clear. I provided my drafts first and used the feedback to improve completeness and clarity, but all original analysis and explanations are my own.

<!-- END QUESTION -->

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **required** to go to the Kernel option and choose **"Restart & Run All"**  before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/ "‌"). **Submissions not prepared correctly as above will lose points.**

Tests for check_bad_variable_names (points: [0.5])
```
>>> from auto_checks import get_default_type_definitions
>>> types_to_check = {'int': int, 'all': all, 'min': min, 'max': max, 'any': any, 'isinstance': isinstance, 'float': float, 'complex': complex, 'str': str, 'bytes': bytes, 'bytearray': bytearray, 'list': list, 'tuple': tuple, 'range': range, 'dict': dict, 'set': set, 'frozenset': frozenset, 'bool': bool, 'NoneType': type(None), 'Ellipsis': type(Ellipsis), 'NotImplemented': type(NotImplemented), 'slice': slice, 'memoryview': memoryview, 'classmethod': classmethod, 'staticmethod': staticmethod, 'property': property, 'object': object}
>>> def are_types_overwritten(types_to_check):
...     """
...     Joseph Scarpa, 02-24-2024
...     
...     Checks if built-in types are overwritten by custom definitions.
...     - types_to_check: A dictionary mapping from type names (as strings) to the custom types defined in the student notebook.
...     
...     Returns:
...     - Boolean indicating whether the built-in types are overridden or not.
...     """
...     built_in_types = get_default_type_definitions()
...     for type_name, type_val in types_to_check.items():
...         if built_in_types.get(type_name) is not type_val:
...             msg = f'NOT YET: It looks like you have created a variable called {type_name}. '
...             msg += 'This is a reserved keyword in Python. Use a different variable name.'
...             print(msg)
...             return False
...     return True
>>> bool(are_types_overwritten(types_to_check))
True
```


Tests for q1 (points: [1, 1, 2])
```
>>> bool(isinstance(q1, tuple) and (isinstance(q1[0], np.float64) or isinstance(q1[0], float)) and (isinstance(q1[1], np.float64) or isinstance(q1[1], float)))
True
```
```
>>> bool(0 < q1[0] < 100 and 0 < q1[1] < 100)
True
```
```
>>> from auto_checks import check_q1
>>> check_q1(q1)
'PASS'
```


Tests for q2_code (points: [1, 1, 2])
```
>>> bool(q2_setting in [1, 2])
True
```
```
>>> bool(isinstance(q2, np.float64) or isinstance(q2, float))
True
```
```
>>> from auto_checks import check_q2
>>> check_q2(q2, q2_setting)
'PASS'
```


Tests for q3 (points: [1, 1, 2, 1, 1, 2])
```
>>> bool(isinstance(q3_1, tuple) and (isinstance(q3_1[0], np.float64) or isinstance(q3_1[0], float)) and (isinstance(q3_1[1], np.float64) or isinstance(q3_1[1], float)))
True
```
```
>>> bool(0 < q3_1[0] < 100 and 0 < q3_1[1] < 100)
True
```
```
>>> from auto_checks import check_q3_1
>>> check_q3_1(q3_1)
'PASS'
```
```
>>> bool(isinstance(q3_2, tuple) and (isinstance(q3_2[0], np.float64) or isinstance(q3_2[0], float)) and (isinstance(q3_2[1], np.float64) or isinstance(q3_2[1], float)))
True
```
```
>>> bool(0 < q3_2[0] < 100 and 0 < q3_2[1] < 100)
True
```
```
>>> from auto_checks import check_q3_2
>>> check_q3_2(q3_2)
'PASS'
```


Tests for q4 (points: [1, 1, 2, 1, 1, 2])
```
>>> bool(isinstance(q4_1, tuple) and (isinstance(q4_1[0], np.float64) or isinstance(q4_1[0], float)) and (isinstance(q4_1[1], np.float64) or isinstance(q4_1[1], float)))
True
```
```
>>> bool(0 < q4_1[0] < 100 and 0 < q4_1[1] < 100)
True
```
```
>>> from auto_checks import check_q4_1
>>> check_q4_1(q4_1)
'PASS'
```
```
>>> bool(isinstance(q4_2, tuple) and (isinstance(q4_2[0], np.float64) or isinstance(q4_2[0], float)) and (isinstance(q4_2[1], np.float64) or isinstance(q4_2[1], float)))
True
```
```
>>> bool(0 <= q4_2[0] <= 5 and 0 <= q4_2[1] <= 5)
True
```
```
>>> from auto_checks import check_q4_2
>>> check_q4_2(q4_2)
'PASS'
```


Tests for q5 (points: [1, 1, 1, 1, 1, 1])
```
>>> import numpy
>>> bool(isinstance(q5_1, numpy.float64) or isinstance(q5_1, float))
True
```
```
>>> from auto_checks import check_q5_1
>>> check_q5_1(q5_1)
'PASS'
```
```
>>> import numpy
>>> bool(isinstance(q5_2, numpy.float64) or isinstance(q5_2, float))
True
```
```
>>> from auto_checks import check_q5_2
>>> check_q5_2(q5_2)
'PASS'
```
```
>>> import numpy
>>> bool(isinstance(q5_3, numpy.float64) or isinstance(q5_3, float))
True
```
```
>>> from auto_checks import check_q5_3
>>> check_q5_3(q5_3)
'PASS'
```
