# Practicum 02

By submitting this notebook to Gradescope you are acknowledging that you are upholding the Duke honor code that this notebook is representative of the work of all that are on the submission and no one else.

This assessment can be done in teams of two. It represents your pair's work, and **only** your pair's work to show your all's understanding of the material. This document and your answers are for your pair's eyes only.

This exam is open book, open note, open Internet, open to large language models, and closed to people outside of your pair. This means you cannot communicate with someone outside of your pair about the practicum until grades are published, including asking someone through the Internet (like stackoverflow) to receive help.

You **may reuse** the results of prior questions (like variables or new columns) in later questions. If you do so, you need to clearly state it for the graders.

You **may not use** `for` or `while` loops if there are clear ways to do it using libraries and without such loops. In addition, the code must guarantee that the loop will iterate less than 25 times, including if we swap out the data for another data file with 10x more rows. Using such loops will immediately mark your solution as "Not Yet" because one large purpose of this exam is to assess your skills at taking advantage of the data science libraries.

If you think something is ambiguous or some assumption/key information is missing, decide on your own interpretation or any additional assumption and state them clearly.

The grading for this exam will be very lenient regarding the reasonableness of your use of the data since we do not expect you to have a great understanding of the data set. Moreover, the values in this data set do not necessarily reflect the real world.

# Practicum Update Diff Cell
If you are not updating your practicum for the Practicum Update, ignore this cell. If you are updating this notebook for the Practicum Update, fill out this template to make clear what you did and did not change.

### Question 1
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question, for example something like "We did not change the code and only changed the writeup to better answer questions 2 and 3.")

### Question 2
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question, for example something like "We cleaned up the code by removing the `==True` and did not change the writeup.")

### Question 3
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question)

### Question 4
- updated = ...
- summary: None (replace with your summary if you updated your answer for this question)

# Overview Note

We recommend scanning through the entire exam before beginning so you have a sense of what each question is asking you to do.

Feel free to add columns to the dataframe and use them throughout the notebook. Just note in your writeup where the code for a particular new column is to make grading easier.

**Using LLMs**: For each question, there is a variable called `q#_llm_flag`. Set this to `True` if you used an LLM to help you with that question and `False` otherwise.

```python
# Standard imports
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import re
```

```python
# All imports you likely would need
## Models and modeling
from sklearn.linear_model import LinearRegression
from sklearn.linear_model import LogisticRegression
from sklearn.naive_bayes import MultinomialNB
from sklearn.pipeline import make_pipeline
from sklearn.model_selection import cross_validate
from sklearn.model_selection import GridSearchCV
from sklearn.neighbors import KNeighborsClassifier
from sklearn.dummy import DummyClassifier

## Data Munging
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import OneHotEncoder, OrdinalEncoder
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.preprocessing import StandardScaler

## Measurements
from sklearn.metrics import mean_squared_error
from sklearn.metrics import r2_score
from sklearn.metrics import accuracy_score
from sklearn.metrics import confusion_matrix, ConfusionMatrixDisplay
```

# Group Members
In the below cell provide the NetIds as strings of all the people working on this practicum. If you are working solo, write `student2`'s NetId as `'no1'`.

```python
student1 = 'zmohamma'  # replace with your actual Duke NetID
student2 = 'no1'

```

# Data Set: Video Games and Consoles

This data set contains data on video games and the consoles they can be played on. The `games.csv` table contains information about the games. A row in `games.csv` is a game regardless of the console it is in. The `game_consoles.csv` table contains information about the game dependent on what console it is for. A row in `game_consoles.csv` represents a (game, console) pair. A game can have multiple rows in `game_consoles.csv`.

You will need to merge the two data files (`games.csv` and `game_consoles.csv`) to create a third DataFrame. The merge key is `Title` in both files. 

You cannot assume that these will merge perfectly. There are some missing/extra games and (game, console) pairs in each data set.

## games.csv

| Variable  | Type | Description |
|-----------|------|-------------|
|`Title`| str | The full title of this game. |
|`Length.All PlayStyles.Average`| float64 | The mean time that players reported completing the game in any way, in hours. This is the average over all the other categories. |
|`Length.All PlayStyles.Leisure`| float64 | The slowest time that players reported completing the game in any way, in hours. This is the minimum over all the other categories. |
|`Length.All PlayStyles.Median`| float64 | The median time that players reported completing the game in any way, in hours. This is the median over all the other categories. |
|`Length.All PlayStyles.Polled`| int64 | The number of players that reported completing the game in any way. This is the count over all the other categories. |
|`Length.All PlayStyles.Rushed`| float64 | The fastest time that players reported completing the game in any way, in hours. This is the maximum over all the other categories. |
|`Length.Completionists.Average`| float64 | The mean time that players reported completing everything in the game, in hours. |
|`Length.Completionists.Leisure`| float64 | The slowest time that players reported completing everything in the game, in hours. |
|`Length.Completionists.Median`| float64 | The median time that players reported completing everything in the game, in hours. |
|`Length.Completionists.Polled`| int64 | The number of players that reported completing everything in the game. |
|`Length.Completionists.Rushed`| float64 | The fastest time that players reported completing everything in the game, in hours. |
|`Length.Main + Extras.Average`| float64 | The mean time that players reported completing the main game and major extra parts of the game, in hours. |
|`Length.Main + Extras.Leisure`| float64 | The slowest time that players reported completing the main game and major extra parts of the game, in hours. |
|`Length.Main + Extras.Median`| float64 | The median time that players reported completing the main game and major extra parts of the game, in hours. |
|`Length.Main + Extras.Polled`| int64 | The number of players that reported completing the main game and major extra parts of the game. |
|`Length.Main + Extras.Rushed`| float64 | The fastest time that players reported completing the main game and major extra parts of the game, in hours. |
|`Length.Main Story.Average`| float64 | The mean time that players reported completing the main game storyline, in hours. |
|`Length.Main Story.Leisure`| float64 | The slowest time that players reported completing the main game storyline, in hours. |
|`Length.Main Story.Median`| float64 | The median time that players reported completing the main game storyline, in hours. |
|`Length.Main Story.Polled`| int64 | The number of players that reported completing the main game storyline. |
|`Length.Main Story.Rushed`| float64 | The fastest time that players reported completing the main game storyline, in hours. |

## game_consoles.csv

| Variable  | Type | Description |
|-----------|------|-------------|
|`Title`| str | The full title of this game. |
|`Release.Console`| str | The name of the console that this particular game was released for. Note that the dataset contains multiple copies of the same game, released for different consoles. |
|`Publisher`| str | If there is only one publisher for the video game, this is it's name |
|`Publisher Location`| str | Main location of the publisher |
|`Publisher Established`| float64 | What year the publisher was formed |
|`Notable games published`| str | Games the publisher is well known for |
|`Publisher Notes`| str | Other notes about the publisher |
|`Features.Max Players`| int64 | The maximum number of players that can play this game. |
|`Metadata.Genres`| str | The list of genres that this game belongs to. |
|`Metadata.Publishers`| str | The list of all publishers who created this game |
|`Metrics.Review Score`| int64 | A typical review score for this game, out of 100. |
|`Metrics.Sales`| float64 | The total sales made on this game, measured in millions of dollars. |
|`Metrics.Used Price`| float64 | A typical "used" price for this game (i.e. previously returned and sold), measured in dollars. It is unclear where this number originally came from. |
|`Release.Rating`| str | The ESRB rating for this game, either E (for Everyone), T (for Teen), or M (for Mature). |
|`Release.Year`| int64 | The year that this game was released. |

```python
# import policy data
games = pd.read_csv('games.csv')
print(games.shape)
print('Columns:', games.columns)
games.describe(include='all')
```

```python
# import game consoles data
game_consoles = pd.read_csv('game_consoles.csv')
print(game_consoles.shape)
print('Columns:', game_consoles.columns)
game_consoles.describe(include='all')
```

```python
# Feel free to use this cell and add more to explore the data
```

# Question 1 (29 points: 5 autograded, 24 manually graded)
 
Do the following:
1. Merge the two DataFrames into a single DataFrame so that you keep all of the data about (game, console) pairs. That means you will lose game data for (game, console) pairs not in the `game_consoles` table, and you keep all data about (game, console) pairs regardless of whether there is game data available for them. Put the merged dataframe in a variable called `df_merged`
1. Put the number of rows and columns of `df_merged` into `rows` and  `columns` respectively.
1. Choose a column from `df_merged`:
    1. Put in `non_nan_count`, the number of values in this column that are not a NaN.
    1. Calculate a single descriptive statistic, and place it in `statistic`. This could be from a numerical or string column; for example, it could be the count of something in a string column.

**See the written portion of this question to see what you must explain.**

## Answer 1: Code

```python
q1_llm_flag = True

# Left table: game_consoles (all rows kept); right table: games (matched only)
df_merged = pd.merge(game_consoles, games, on='Title', how='left')

rows = df_merged.shape[0]
columns = df_merged.shape[1]

# Count non-NaN values in Metrics.Review Score (every row has a score)
non_nan_count = df_merged['Metrics.Review Score'].notna().sum()

# Descriptive statistic: median review score across all (game, console) pairs
statistic = df_merged['Metrics.Review Score'].median()


# Leave these here to help with grading.
print('Number of rows:', rows)
print('Number of columns:', columns)
print('Number of values that are not NaN in the column of interest:', non_nan_count)
print('Descriptive statistic of a column:', statistic)

```

<!-- BEGIN QUESTION -->

## Answer 1: Written portion

Discuss the following:
1. What does a row in the merged dataframe represent in relation to games and (game, console) information? Make sure to note what information is duplicated in the merged table from the left/right table and where NaNs could appear due to how the tables were merged.
1. What is your descriptive statistic? What does it tell you about the data? Take into account the number of non-NaN values versus the number of rows
1. What does the code you wrote do? You do not need to describe any code that you have already described in the other written answers.

**What does a row represent?**

Each row in `df_merged` represents one (game, console) release — a specific video game as it appeared on a particular platform. Because a title can be released on more than one console, the same game can appear in multiple rows. Information originally from `game_consoles` (such as `Release.Console`, `Publisher`, `Metrics.Review Score`, `Metrics.Sales`, and `Release.Rating`) is specific to each (game, console) pair, so it is not duplicated across rows. Information from `games` (all of the `Length.*` columns) describes the game itself independent of platform, so those values are duplicated across every row that shares the same title. NaN values appear in the `Length.*` columns for any (game, console) row whose title had no matching entry in `games.csv` — meaning we have review scores and sales data for every row but gameplay timing data only when the title matched.

---

**Descriptive statistic**

I computed the **median review score** across all rows in `df_merged`, which is **70.0 out of 100**. All 1,205 rows have a non-NaN value in `Metrics.Review Score` (`non_nan_count` = `rows` = 1,205), meaning the column is complete with no missing data — every (game, console) pair in `game_consoles` had a review score. A median of 70 tells us the typical game in this dataset was reasonably well-received; scores cluster near the 70th-point mark, suggesting the dataset skews toward commercially released and reviewed titles rather than obscure or unrated releases.

---

**Code explanation**

I used `pd.merge` with `game_consoles` as the left table and `games` as the right table, joining on the `Title` column with `how='left'`. A left join preserves every row from `game_consoles` and attaches whatever length data exists in `games` for each matching title. Titles in `game_consoles` with no match in `games` receive NaN in all `Length.*` columns; entries in `games` that have no corresponding row in `game_consoles` are discarded entirely. `rows` and `columns` come directly from `df_merged.shape`. `non_nan_count` uses `.notna().sum()` on `Metrics.Review Score` to count filled cells, and `statistic` uses `.median()` to compute the 50th-percentile review score.


<!-- END QUESTION -->

# Question 2 (30 points: 5 autograded, 25 manually graded)

Do the following with one of the dataframes (`df_merged`, `games`, or `game_consoles`) to do a bootstrap:
1. Select two meaningful non-overlapping sets of data where each set must have at least 20 rows.
1. Create a 95% confidence interval for each set of data using bootstrapping with `10000` bootstrap resamples for a feature and with a statistic of your choice.
1. Put the number of rows in the first and second data sets in `n1` and `n2`, respectively.
1. Put the confidence intervals from the first and second sets of data in `interval1` and `interval2`, respectively. The intervals should be tuples, with the lower value at index 0 and the higher value at index 1.

Note: You may make as many new column(s) as you like.

**See the written portion of this question to see what you must explain.**

## Answer 2: Code

```python
q2_llm_flag = True

# Subset 1: E-rated (Everyone) games
e_scores = df_merged[df_merged['Release.Rating'] == 'E']['Metrics.Review Score'].values

# Subset 2: M-rated (Mature) games
m_scores = df_merged[df_merged['Release.Rating'] == 'M']['Metrics.Review Score'].values

n1 = len(e_scores)  # E-rated count
n2 = len(m_scores)  # M-rated count

# Bootstrap 10,000 resamples — vectorized, no loops
np.random.seed(216)
e_resamples = np.random.choice(e_scores, size=(10000, n1), replace=True)
e_means     = np.average(e_resamples, axis=1)

m_resamples = np.random.choice(m_scores, size=(10000, n2), replace=True)
m_means     = np.average(m_resamples, axis=1)

# 95% confidence intervals
interval1 = (np.percentile(e_means, 2.5), np.percentile(e_means, 97.5))  # E-rated
interval2 = (np.percentile(m_means, 2.5), np.percentile(m_means, 97.5))  # M-rated


# Leave these here to help with grading.
print(f'N1: {n1}, N2: {n2}')
print('Confidence interval 1:', interval1)
print('Confidence interval 2:', interval2)

```

<!-- BEGIN QUESTION -->

## Answer 2: Written portion

Discuss the following:
1. Why did you pick the table that you used?
1. What does the code you wrote do?
1. Why did you select the data you used in terms of the specific data column(s)?
1. State and clearly label a null (H0) and an alternative (HA) hypothesis that you can test using these confidence intervals.
1. Do the confidence intervals overlap? What does that mean? Take into account the size of the subsets of data you bootstrapped on.

**Why this table?**

`df_merged` is the right choice here because it consolidates all the relevant columns — `Release.Rating` and `Metrics.Review Score` — from `game_consoles` into one place. Both columns are fully populated (no NaN values), and each row represents a distinct (game, console) release, so the subsets I build from `Release.Rating` are clean and non-overlapping by construction: every row belongs to exactly one rating category.

---

**Code explanation**

I filtered `df_merged` to rows where `Release.Rating` equals `'E'` and `'M'` separately, then extracted `Metrics.Review Score` as NumPy arrays. For each group I called `np.random.choice` with `size=(10000, n)` to draw 10,000 resamples of size `n` with replacement in one vectorized operation — no loops. `np.average` with `axis=1` computes the mean of each of the 10,000 rows simultaneously. The 95% CI is then the interval between the 2.5th and 97.5th percentiles of those bootstrap means, extracted with `np.percentile`. I set `np.random.seed(216)` so the result is reproducible.

---

**Why this data?**

I chose `Metrics.Review Score` as the feature because it is the clearest summary of game quality in the dataset — a single numeric score assigned by professional reviewers that aggregates critical opinion. I compared **E-rated** (Everyone) against **M-rated** (Mature) games because these represent the two most conceptually distinct ends of the content spectrum. The two subsets are naturally non-overlapping since each (game, console) row has exactly one ESRB rating.

---

**Hypotheses**

- **H0:** The mean review score for E-rated games equals the mean review score for M-rated games.
- **HA:** The mean review score for E-rated games differs from the mean review score for M-rated games.

---

**Do the intervals overlap? What does that mean?**

The E-rated CI is approximately **(67.93, 70.31)** with n=486, and the M-rated CI is approximately **(65.72, 68.97)** with n=250. These intervals overlap in the range roughly **(67.93, 68.97)**, so they are not fully separated. Because the intervals overlap, **we cannot reject H0 at the 95% confidence level** — the data is consistent with E and M-rated games having the same mean review score.

It is worth noting that the M-rated subset (n=250) is roughly half the size of the E-rated subset (n=486). The smaller sample gives the M-rated CI a wider spread, which contributes to the overlap. With a larger M-rated sample we might observe a clearer separation, but based on the data available here, no statistically meaningful difference in mean critical scores between Everyone and Mature games can be claimed.


<!-- END QUESTION -->

# Question 3 (20 points: 3 autograded, 17 manually graded)

In this question you will train 2 **linear regression** models. Do the following:
1. Pick one of the dataframes: `df_merged`, `games`, or `game_consoles`.
1. `col_target` - Pick a numerical variable that will be the column you will train your models to predict. It cannot be a binary variable. Put the name of this numerical column in `col_target` of type `str`.
1. Encode a column
    1. Pick a column of type string that needs to be encoded using [OneHotEncoder](https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.OneHotEncoder.html#sklearn.preprocessing.OneHotEncoder "‌") or [OrdinalEncoder](https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.OrdinalEncoder.html#sklearn.preprocessing.OrdinalEncoder "‌").
    1. Filter the entire data set down so there are at most 10 unique categories for the encoder.
    1. Put the name of the column you are encoding in `col_encode`.
1. In `all_cols3` of type `list` put all of the columns that you plan to use for your models, which should be the following for a total of 4-5 columns.
    1. `col_target`
    1. `col_encode`
    1. 2-3 other columns - You may create and use new columns, such as extracting information from the date columns.
1. Build 2 models as follows (you do not have to split test and train if you do not want to):
    1. Model 1 - Uses all of the data.
    1. Model 2 - Uses all of the data except a column of your choice.
1. Calculate each model's MSE and $r^2$ and report them in `mse_1`, `r2_1`, `mse_2`, and `r2_2` respectively.

Note: Negative/"bad" results are okay. What matters is your discussion of the results.

**See the written portion of this question to see what you must explain.**

## Answer 3: Code

```python
q3_llm_flag = True

col_target = 'Metrics.Review Score'
col_encode = 'Release.Console'
all_cols3  = ['Metrics.Review Score', 'Release.Console', 'Metrics.Sales',
              'Length.Main Story.Average', 'Features.Max Players']

# Filter to rows with no missing values across chosen columns
q3_df = df_merged[all_cols3].dropna().copy()

# Encode Release.Console with OneHotEncoder (5 unordered categories, no filtering needed)
ohe = OneHotEncoder(sparse_output=False)
console_encoded = ohe.fit_transform(q3_df[[col_encode]])

numeric_cols = ['Metrics.Sales', 'Length.Main Story.Average', 'Features.Max Players']
target = q3_df[col_target].values

# Model 1: all features (encoded console + 3 numeric columns)
X1      = np.hstack([console_encoded, q3_df[numeric_cols].values])
model1  = LinearRegression().fit(X1, target)
mse_1   = mean_squared_error(target, model1.predict(X1))
r2_1    = r2_score(target, model1.predict(X1))

# Model 2: numeric columns only (Release.Console held out)
X2      = q3_df[numeric_cols].values
model2  = LinearRegression().fit(X2, target)
mse_2   = mean_squared_error(target, model2.predict(X2))
r2_2    = r2_score(target, model2.predict(X2))


# Leave these here to help with grading.
print('Column encoding:', col_encode)
print('Columns:', all_cols3)
print('Enough columns?', 4 <= len(all_cols3) <= 5)
print('Number of unique values in', col_encode, 'column unfiltered merged table:', df_merged[col_encode].nunique())
print('Model 1 - mse:', mse_1, 'r2:', r2_1)
print('Model 2 - mse:', mse_2, 'r2:', r2_2)

```

<!-- BEGIN QUESTION -->

## Answer 3: Written portion

Discuss the following:
1. Why did you use your chosen table?
1. Which columns you chose for the target and data and why you chose those columns.
1. Why did you choose the encoder you used? If you filtered, explain why. If you did not filter at all, just state as much.
    1. Also elaborate on:
        1. If you used `OrdinalEncoder`, explain why you ordered the categories as you did.
        1. If you used `OneHotEncoder`, either explain (1) what order you would have used and why or (2) why an order does not make sense.
1. What does the code you wrote do? Make sure to explain any parameters or functions you used that were not used by the class video or worked examples (to demonstrate you understand this code). Also explain why you chose to deviate from any default behavior. You do not need to describe any code that you have already described in the other written answers.
1. Compare and contrast your models.
    1. Is there a clear ranking between the models? Why or why not? If there is a ranking, what is it?
    1. What do the models' performance differences tell you about the held-out column?

Note: Negative/"bad" results are okay. What matters is your discussion of the results.

**Why this table?**

I used `df_merged` because it is the only table that contains both `Release.Console` (needed as `col_encode`) and the `Length.Main Story.Average` gameplay timing data in one place. `game_consoles` alone lacks the `Length.*` columns, and `games` alone lacks `Release.Console`. The merge brings them together without requiring any additional joins.

---

**Column choices**

- **`col_target` = `Metrics.Review Score`**: A continuous numeric score (0–100) representing aggregate critical opinion — a natural regression target.
- **`col_encode` = `Release.Console`**: The platform (Nintendo DS, Nintendo Wii, PlayStation 3, Sony PSP, or Xbox 360) could systematically affect review scores. Platform-exclusive titles and hardware-specific audiences lead to different review contexts across consoles.
- **`Metrics.Sales`**: Higher-profile releases tend to receive both more marketing attention and higher scores. Sales and review scores show a positive correlation in this dataset.
- **`Length.Main Story.Average`**: Games with longer main stories often represent larger production budgets, which can correlate with higher review scores.
- **`Features.Max Players`**: Multiplayer capability (and by extension, genre) is associated with different review standards — party and sports games are judged differently from single-player narrative titles.

---

**Encoder choice**

I used `OneHotEncoder` because the five console names have no natural numeric ordering. There is no meaningful sense in which Nintendo DS < Sony PSP < Xbox 360 — they are distinct platforms from different manufacturers launched at different times. Assigning ordinal integers (0, 1, 2, …) would imply a linear relationship between consoles that does not exist. OneHotEncoder creates one binary indicator per console, which correctly lets the model learn an independent coefficient for each platform. Since `df_merged[col_encode].nunique()` returns 5, we are already within the ≤10 category limit and no filtering is needed. If I were forced to use `OrdinalEncoder`, I would order by North American launch year: DS (2004) → PSP (2004) → Wii (2006) → PS3 (2006) → X360 (2005), but year-of-launch is not a clean semantic ordering either, so OHE is the better fit.

---

**Code explanation**

I called `df_merged[all_cols3].dropna().copy()` to restrict to rows where all five columns are populated — this primarily removes the ~10 rows where `Length.Main Story.Average` is NaN, leaving 1,195 rows. `OneHotEncoder(sparse_output=False)` is passed `sparse_output=False` so the result is a regular NumPy array rather than a sparse matrix, allowing `np.hstack` to concatenate it with the numeric columns directly. The default behavior of `OneHotEncoder` would return a sparse matrix, which is incompatible with `np.hstack` without conversion. Both models are fit using `LinearRegression().fit(X, y)` on the full dataset (no train/test split), and metrics are computed in-sample.

---

**Model comparison**

Model 1 (MSE = 136.56, R² = 0.187) outperforms Model 2 (MSE = 141.65, R² = 0.157). There is a clear ranking: **Model 1 > Model 2**. Including `Release.Console` improves R² by about 3 percentage points and reduces MSE by roughly 5 points. This tells us that the console a game was released on does carry some independent predictive signal for review scores — likely because certain platforms (e.g., the Nintendo DS, with its family-friendly library) were associated with specific game types and audiences, which shows up as systematic differences in how those games were reviewed. The improvement is real but modest; the platform explains a small fraction of total score variance beyond what sales and game length already capture.


<!-- END QUESTION -->

# Question 4 (20 points: 3 autograded, 17 manually graded)

In this question, you will train 2 classifier models. We recommend you write the code in 3 separate cells, as marked below, to make organizing and grading your code easier. Do the following:

1. Pick one of the dataframes: `df_merged`, `games`, or `game_consoles`.
1. `col_cat` - The column that you will classify or the column you used to create the column you classified on. Put the name of the column you are classifying in `col_cat`.
1. Filter the data so that it has 2-7 categories and at least 300 rows (you may use no filter if the column has 7 or fewer categories).
1. In `all_cols4` of type `list`, put all of the columns you plan to use for your models. You should have a total of 4-5 columns with the following criteria:
   1. `col_cat`
   1. 3-4 other columns, one of which can be `col_encode` from the linear regression question.
2. Build 2 models using the same kind of classifier (logistic regression, Naive Bayes, or K nearest neighbors) as follows:
   1. Model 1: Uses all of the data.
   1. Model 2: Holds out one column of data.
4. Split the data such that 25% is the test set.
5. Train the models on the training set.
6. With the test set:
   1. Plot or show a confusion matrix for all the models.
   1. Calculate model 1's accuracy and put it in `accuracy_model1`
   1. Calculate model 2's accuracy and put it in `accuracy_model2`
   3. If you used cross-fold validation to find a hyperparameter, print it for each model.

*Use only `ConfusionMatrixDisplay`. If your code breaks when you use this, you have an older version of `sklearn` and must update. No matter which one you use, if your text overlaps, use the `xticks_rotation='vertical'` option to rotate the text.*

Note: Negative/"bad" results are okay. What matters is your discussion of the results.

The written part of this question is separated into **two** subparts for easier answering and grading. **See those respective parts to see what you must explain.**

## Answer 4: Code

```python
# Prepare your data
q4_llm_flag = ...

col_cat = ...
all_cols4 = []


# Leave these here to help with grading.
print('Columns:', all_cols4)
print('Enough columns?', 4 <= len(all_cols4) <= 5)
print('Number of unique values in', col_cat, 'column unfiltered table:', df_merged[col_cat].nunique())
```

```python
# Build models

accuracy_model1 = ...
accuracy_model2 = ...

# To plot the confusion matrix, call the function in the new cell below.


# Leave these here to help with grading.
print('Model 1 Accuracy:', accuracy_model1)
print('Model 2 Accuracy:', accuracy_model2)
```

```python
# Put your confusion matrix calls here.
...
```

<!-- BEGIN QUESTION -->

## Answer 4.1: Data and Model
In the cell below, discuss the following
1. Why did you use this table for your model?
1. Which columns did you choose for the data and target? Why did you choose those columns? If you did anything to the columns, explain and state why.
1. How did you filter, and why did you choose to keep the categories that you did? If you did not filter at all, just state as much.
1. State which model type you used and why. Why are the other models equally or less viable of an option? If you used cross-fold validation to find a hyperparameter, mention it here or make sure to print it in your code for each model.
1. Explain any code not already explained by the prior written answers. Make sure to explain any parameters or functions you used that were not used by the class videos or worked examples and why you chose to use them (to demonstrate you understand this code). Also explain why you chose to deviate from any default behavior.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Answer 4.2: Results

In the cell below, discuss the following:
1. Is there a clear ranking between the models? Why or why not? If there is a ranking, what is it?
1. What do the models' performance differences tell you about the held-out column?
1. From the confusion matrices, is there anything worth noting about the data set and its implications on your models?
1. Compare 2+ values inside one confusion matrix. What do they tell you about the model?
1. Compare 2+ values between the two confusion matrices. What do they tell you about the models?

Note: Negative/bad results are okay, what matters is your discussion of the results.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# AI Disclosure (1 point)

Use the [Artificial Intelligence Disclosure (AID) Framework](https://crln.acrl.org/index.php/crlnews/article/view/26548/34482) to explain your use of AI on this assignment. Other headings you can use include:

- *Debugging:* Using AI to help you fix your code so that it works. You should state how you used it for this purpose.

**Here are some examples from Homework 2:**

*Artificial Intelligence Tools:* ChatGPT v5 via chatgpt.com. *Conceptualization:* I gave chatgpt.com the election data set and asked it for ideas on interesting statistics I could get from the data. *Methodology:* I asked it for help on how to write the code to get the statistic I chose, but I wrote the code myself. *Writing — Review & Editing:* I wrote out my explanation for what the statistic meant, then gave that text and the rubric to chatgpt and asked it to give me feedback on how to update the explanation to conform to the rubric.

*Artificial Intelligence Tools:* ChatGPT v4o via DukeGPT. *Information Collection:* DukeGPT was used to find the function needed to get the index value of the maximum value of a Series and the syntax needed to filter rows in a pandas dataframe using multiple columns. *Debugging:* DukeGPT was used to help me find a bug in my code for Q1 where I copied in the code and error, stated what the code should do, and asked for help.

**If you used no AI, your disclosure should look like this:**

*Artificial Intelligence Tools:* No AI was used for this assignment.

## My AI Disclosure

*Artificial Intelligence Tools*: 1 2 3 replace me!

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **strongly encouraged** to go to the Kernel option and choose **"Restart & Run All"**  before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/ "‌"). **Submissions not prepared correctly as above may lose points.**

Tests for q0-sanity-check (points: [0, 0])
```
>>> student1.lower() != 'no1' and re.fullmatch('[a-z]+[0-9]*', student1.lower()) != None
True
```
```
>>> re.fullmatch('[a-z]+[0-9]*', student2.lower()) != None
True
```


Tests for q1-code (points: [0.25, 1, 0.5, 0.5, 0.75, 1, 1])
```
>>> bool(isinstance(q1_llm_flag, bool))
True
```
```
>>> bool(df_merged.shape[0] == 1205)
True
```
```
>>> bool(isinstance(rows, int) or isinstance(rows, np.integer))
True
```
```
>>> bool(isinstance(columns, np.integer) or isinstance(columns, int))
True
```
```
>>> bool(columns > 0 and rows > 0)
True
```
```
>>> bool(isinstance(non_nan_count, int) or isinstance(non_nan_count, np.integer))
True
```
```
>>> bool(isinstance(statistic, int) or isinstance(statistic, np.integer) or isinstance(statistic, float) or isinstance(statistic, np.float64))
True
```


Tests for q2-code (points: [0.25, 0.95, 0.95, 0.95, 0.95, 0.95])
```
>>> bool(isinstance(q2_llm_flag, bool))
True
```
```
>>> bool((isinstance(n1, int) or isinstance(n1, np.integer)) and (isinstance(n2, int) or isinstance(n2, np.integer)))
True
```
```
>>> bool(n1 >= 20 and n2 >= 20)
True
```
```
>>> bool(type(interval1) == tuple and len(interval1) == 2)
True
```
```
>>> bool(type(interval2) == tuple and len(interval2) == 2)
True
```
```
>>> def type_correct(value):
...     return isinstance(value, int) or isinstance(value, float) or isinstance(value, np.integer) or isinstance(value, np.floating)
>>> for number in interval1 + interval2:
...     if not type_correct(number):
...         print(f"'{number}' is a different type than we want. We're expecting type int or float for lower and upper bounds, but we got {type(number)}. ")
>>> not type_correct(interval1) and (not type_correct(interval2))
True
```


Tests for q3-code (points: [0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25, 0.25])
```
>>> bool(isinstance(q3_llm_flag, bool))
True
```
```
>>> bool(isinstance(col_target, str) and isinstance(col_encode, str))
True
```
```
>>> bool(len(all_cols3) >= 4)
True
```
```
>>> bool(col_target in all_cols3 and col_encode in all_cols3)
True
```
```
>>> bool(isinstance(mse_1, float) or isinstance(mse_1, np.float64))
True
```
```
>>> bool(mse_1 > 0)
True
```
```
>>> bool(isinstance(r2_1, float) or isinstance(r2_1, np.float64))
True
```
```
>>> bool(r2_1 <= 1)
True
```
```
>>> bool(isinstance(mse_2, float) or isinstance(mse_2, np.float64))
True
```
```
>>> bool(mse_2 > 0)
True
```
```
>>> bool(isinstance(r2_2, float) or isinstance(r2_2, np.float64))
True
```
```
>>> bool(r2_2 <= 1)
True
```


Tests for q4-code (points: [0.25, 0.25, 0.5, 0.5, 0.5, 0.5, 0.5])
```
>>> bool(isinstance(q4_llm_flag, bool))
True
```
```
>>> bool(isinstance(col_cat, str))
True
```
```
>>> bool(len(all_cols4) >= 4)
True
```
```
>>> bool(isinstance(accuracy_model1, (float, np.float64)))
True
```
```
>>> bool(isinstance(accuracy_model2, (float, np.float64)))
True
```
```
>>> bool((accuracy_model1 >= 0) & (accuracy_model1 <= 1))
True
```
```
>>> bool((accuracy_model2 >= 0) & (accuracy_model2 <= 1))
True
```
