# Homework 09
In this assignment, you will practice writing SQL queries to pull data from a relational database. 

## The congress database
A visual schema of the congress database is also available as a pdf along with this file - we encourage you to start by taking a look to get familiar. This database contains the history of the members of the United States congress through the 115th congress (the data end by 2019) as well as a good deal of voting data from 2015-2016. The visual schema shows each table in the database as a yellow box with the table name at the top and the column names listed below. The arrows showing which keys/identifiers match between different tables for join operations. At the bottom you can see previews of the tables. One in particular, the `cur_members` table, contains data about members of the 115th congress (note - everywhere you see us refer to "current" members below, we are referring to those in the `cur_members` table). Some of these members are still serving as of the 118th congress, which began in January 2023.

## How to use SQL and report your results

You should use Python sqlite3 and Pandas to run your SQL queries in this assignment. We recommend using multiline strings for your queries for readability; remember that multiline strings in Python are enclosed by triple quotes for example,

```
my_multiline_string = """Hi 
There!"""
```

See the worked examples for an example query and it's result.

## Questions

```python
# Run but do not modify the following code
# to import sqlite3 and pandas, and to connect
# to the congress database
import sqlite3
import pandas as pd
import numpy as np
conn = sqlite3.connect("congress")
# If you cannot load the database and see error message like "no such table", try use absolute path like following:
# import sqlite3
# import pandas as pd
# import numpy as np
# import os.path
# BASE_DIR = os.path.dirname(os.path.abspath("congress"))
# db_path = os.path.join(BASE_DIR, "congress")
# conn = sqlite3.connect(db_path)
```

### Question 1 (4 points)
Who are all of the members of congress from North Carolina (`NC`) in the `cur_members` table? Show all of the information about them from the `cur_members` table. Put your resulting DataFrame in `q1`.

```python
# Start with a SQL query to extract the information
query1 = """
SELECT *
FROM cur_members
WHERE state = 'NC'
"""

# Execute queries using pandas
q1 = pd.read_sql(query1, conn)


# Leave this line here so that the table is shown
q1
```

### Question 2 (4 points)
List the five youngest female (`gender = 'F'`) members of the Congress from the `cur_members` table in descending order of `birthday`. Show all information about them from the `cur_members` table. Put your resulting DataFrame in `q2`.

```python
# Start with a SQL query to extract the information
query2 = """
SELECT *
FROM cur_members
WHERE gender = 'F'
ORDER BY birthday DESC
LIMIT 5
"""

# Execute queries using pandas
q2 = pd.read_sql(query2, conn)


# Leave this line here so that the table is shown
q2
```

### Question 3 (4 points)
For the current members of congress, show the average, minimum, and maximum age (in years) of the representatives (i.e., `type = 'rep'`). Show the same for senators (i.e., `type = 'sen'`). Compute ages in years as of the current date (`current_date`). You may compute the age of an individual as simply `current_date - birthday`. The autograder will accept answers computed during any day in the semester.

Note that this is not *exactly* correct due to date rounding error. A more precise calculation of age in years using [sqlite date and time functions](https://sqlite.org/lang_datefunc.html) looks like `(julianday(current_date)-julianday(birthday)) / 365.25`. We will accept solutions using either, you will get very similar results.

Put your resulting DataFrame in `q3` with colums of `type`, `avg`, `min`, and `max`.

```python
# Start with a SQL query to extract the information
query3 = """
SELECT
    type,
    AVG((julianday(current_date) - julianday(birthday)) / 365.25) AS avg,
    MIN((julianday(current_date) - julianday(birthday)) / 365.25) AS min,
    MAX((julianday(current_date) - julianday(birthday)) / 365.25) AS max
FROM cur_members
GROUP BY type
"""

# Execute queries using pandas & check for the empty string
q3 = pd.read_sql(query3, conn)



# Leave this line here so that the table is shown
q3
```

### Question 4 (4 points)
Which parties (`party`) have had at least 50 senators (`sen`) past and present? For each, show the `party` and the number of senators (`sen`) they have had. 

Hint: Be sure not to double count persons who served multiple terms in the senate. You can use `COUNT(DISTINCT column)` to get the number of distinct values on given `column`.

Put your resulting DataFrame in `q4` with colums of `party` and `num_senators`.

```python
# Start with a SQL query to extract the information
query4 = """
SELECT party, COUNT(DISTINCT person_id) AS num_senators
FROM person_roles
WHERE type = 'sen'
GROUP BY party
HAVING COUNT(DISTINCT person_id) >= 50
"""

# Execute queries using pandas
q4 = pd.read_sql(query4, conn)


# Leave this line here so that the table is shown
q4
```

### Question 5 (4 points)
Among the past and present members members of congress who cast at least one vote in the `person_votes` table (i.e., the members whose `person_id`s appear at all in the `person_votes` table), which five members cast the fewest votes? Show their `first_name`s, `last_name`s, and the number of votes they cast in the `person_votes` table. Put your resulting DataFrame in `q5` with colums of `first_name`, `last_name`, and `num_votes`.

Note: There is a vote called "Not Voting", you do not need to treat them specially. Or in other words, do not filter them out when counting, count them as a vote.

```python
# Start with a SQL query to extract the information
query5 = """
SELECT p.first_name, p.last_name, COUNT(pv.vote_id) AS num_votes
FROM persons p
JOIN person_votes pv ON p.id = pv.person_id
GROUP BY p.id, p.first_name, p.last_name
ORDER BY num_votes ASC
LIMIT 5
"""

# Execute queries using pandas
q5 = pd.read_sql(query5, conn)


# Leave this line here so that the table is shown
q5
```

### Question 6 (4 points)
In this question, we would like to consider the change in the gender representation of the US House of Representatives over time. Find all of the `start_date`s from the `person_roles` table in which at least 40 women began terms as representatives in the house (i.e., `type='rep'`). Show the `start_date`s and the number of women beginning terms on those dates. Order the results by `start_date` from oldest to most recent. Put your resulting DataFrame in `q6` with colums of `start_date` and `house_women`.

```python
# Start with a SQL query to extract the information
query6 = """
SELECT pr.start_date, COUNT(DISTINCT p.id) AS house_women
FROM person_roles pr
JOIN persons p ON pr.person_id = p.id
WHERE pr.type = 'rep'
  AND p.gender = 'F'
GROUP BY pr.start_date
HAVING COUNT(DISTINCT p.id) >= 40
ORDER BY pr.start_date ASC
"""

# Execute queries using pandas
q6 = pd.read_sql(query6, conn)


# Leave this line here so that the table is shown
q6
```

<!-- BEGIN QUESTION -->

# AI Disclosure

Use the **Artificial Intelligence Disclosure (AID) Framework** to explain your use of AI on this assignment. Other headings you can use include:

- *Debugging:* Using AI to help you fix your code so that it works. You should state how you used it for this purpose.

Here are some examples:

*Artificial Intelligence Tools:* ChatGPT v5 via chatgpt.com. *Conceptualization:* I gave chatgpt.com the election data set and asked it for ideas on interesting statistics I could get from the data. *Methodology:* I asked it for help on how to write the code to get the statistic I chose, but I wrote the code myself. *Writing — Review & Editing:* I wrote out my explanation for what the statistic meant, then gave that text and the rubric to chatgpt and asked it to give me feedback on how to update the explanation to conform to the rubric.

*Artificial Intelligence Tools:* ChatGPT v4o via DukeGPT. *Information Collection:* DukeGPT was used to find the function needed to get the index value of the maximum value of a Series and the syntax needed to filter rows in a pandas dataframe using multiple columns. *Debugging:* DukeGPT was used to help me find a bug in my code for Q1 where I copied in the code and error, stated what the code should do, and asked for help.

## My AI Disclosure

*Artificial Intelligence Tools:* Claude (Anthropic) via claude.ai.

*Debugging:* Claude was used to help identify a syntax issue in one of the SQL queries — specifically, I copied in the error message, described what the query was supposed to do, and asked for help resolving the issue. All queries were written and understood by me; Claude was only consulted to diagnose the specific error.

*Grammar and Formatting:* I also used Claude (Anthropic) to help fix minor grammar and formatting issues in my writeup. I pasted sections of text into Claude to check for clarity or proper markdown formatting, but all ideas and substantive content are my own; Claude was only used to correct language and formatting mistakes.

<!-- END QUESTION -->

## Submitting

You should make sure any code that you write to answer the questions is included in this notebook. You are **required** to go to the Kernel option and choose **"Restart & Run All"**  before submission. Double check that your entire notebook runs correctly and generates the expected output. Finally, make sure to save your work (timestamp at the top tells you the last checkpoint and whether there are unsaved changes). When you finish, submit your assignment at [Gradescope](http://gradescope.com/ "‌"). **Submissions not prepared correctly as above will lose points.**

Tests for check_bad_variable_names (points: [0.5])
```
>>> from auto_checks import get_default_type_definitions
>>> types_to_check = {'int': int, 'all': all, 'min': min, 'max': max, 'any': any, 'isinstance': isinstance, 'float': float, 'complex': complex, 'str': str, 'bytes': bytes, 'bytearray': bytearray, 'list': list, 'tuple': tuple, 'range': range, 'dict': dict, 'set': set, 'frozenset': frozenset, 'bool': bool, 'NoneType': type(None), 'Ellipsis': type(Ellipsis), 'NotImplemented': type(NotImplemented), 'slice': slice, 'memoryview': memoryview, 'classmethod': classmethod, 'staticmethod': staticmethod, 'property': property, 'object': object}
>>> def are_types_overwritten(types_to_check):
...     """
...     Joseph Scarpa, 02-24-2024
...     
...     Checks if built-in types are overwritten by custom definitions.
...     - types_to_check: A dictionary mapping from type names (as strings) to the custom types defined in the student notebook.
...     
...     Returns:
...     - Boolean indicating whether the built-in types are overridden or not.
...     """
...     built_in_types = get_default_type_definitions()
...     for type_name, type_val in types_to_check.items():
...         if built_in_types.get(type_name) is not type_val:
...             msg = f'NOT YET: It looks like you have created a variable called {type_name}. '
...             msg += 'This is a reserved keyword in Python. Use a different variable name.'
...             print(msg)
...             return False
...     return True
>>> are_types_overwritten(types_to_check)
True
```


Tests for q1 (points: [1, 1, 1, 1])
```
>>> q1.shape
(15, 9)
```
```
>>> print(sorted(q1.columns))
['birthday', 'first_name', 'gender', 'id', 'last_name', 'party', 'religion', 'state', 'type']
```
```
>>> from auto_checks import check_q1_state
>>> check_q1_state(q1)
'PASS'
```
```
>>> from auto_checks import check_q1
>>> check_q1(q1)
'PASS'
```


Tests for q2 (points: [1, 1, 1, 1])
```
>>> q2.shape
(5, 9)
```
```
>>> print(sorted(q2.columns))
['birthday', 'first_name', 'gender', 'id', 'last_name', 'party', 'religion', 'state', 'type']
```
```
>>> from auto_checks import check_q2_gender
>>> check_q2_gender(q2)
'PASS'
```
```
>>> from auto_checks import check_q2
>>> check_q2(q2)
'PASS'
```


Tests for q3 (points: [1, 1, 0.5, 0.5, 0.5, 0.5])
```
>>> q3.shape
(2, 4)
```
```
>>> print(sorted(q3.columns))
['avg', 'max', 'min', 'type']
```
```
>>> from auto_checks import check_q3_type
>>> check_q3_type(q3)
'PASS'
```
```
>>> from auto_checks import check_q3_avg
>>> check_q3_avg(q3)
'PASS'
```
```
>>> from auto_checks import check_q3_min
>>> check_q3_min(q3)
'PASS'
```
```
>>> from auto_checks import check_q3_max
>>> check_q3_max(q3)
'PASS'
```


Tests for q4 (points: [1, 1, 1, 1])
```
>>> q4.shape
(5, 2)
```
```
>>> print(sorted(q4.columns))
['num_senators', 'party']
```
```
>>> from auto_checks import check_q4_party
>>> check_q4_party(q4)
'PASS'
```
```
>>> from auto_checks import check_q4
>>> check_q4(q4)
'PASS'
```


Tests for q5 (points: [1, 1, 2])
```
>>> q5.shape
(5, 3)
```
```
>>> print(sorted(q5.columns))
['first_name', 'last_name', 'num_votes']
```
```
>>> from auto_checks import check_q5
>>> check_q5(q5)
'PASS'
```


Tests for q6 (points: [1, 1, 2])
```
>>> q6.shape
(13, 2)
```
```
>>> print(sorted(q6.columns))
['house_women', 'start_date']
```
```
>>> from auto_checks import check_q6
>>> check_q6(q6)
'PASS'
```
