# Worked Example 02

In this exercise, you will practice selecting and transforming data using the `numpy` and `pandas` Python libraries. You can search on the web for documentation and techniques to help you solve the problem. See especially the official [Numpy documentation](https://numpy.org/doc/stable/) and [Pandas documentation](https://pandas.pydata.org/docs/user_guide/index.html).

For every sub-question under every Question, please save results in corresponding variables. For instance, results of Question 1 should be saved in variables named `q1_1`, `q1_2`, ... , `q1_10`. 

Unlike Homeworks, the autograders for worked examples are embedded with the questions so that you can run the tests yourself. Read the instructions below carefully.

### Autograder Instructions

Run the following cell to install the `Otter` autograder package the first time it is run. Afterwards, running it again will print out its version number if `Otter` is correctly installed.

```python
try:
    import otter
    grader = otter.Notebook("worked-example02.ipynb")
    print(otter.__version__)
    if (otter.__version__ != '6.1.6'): #update for latest otter version
        !pip install -q -U --user otter-grader
except Exception as e:
    !pip install -q otter-grader
```

To grade your own work, simply run the cell starting with `grader.check` immediately after each question. The cell calls `Otter` to run the tests for all subquestions and generate a report of what test(s) you pass/do not pass in the same style as on Gradescope.

## Part 1: Numpy
In this part you will practice with the functions and conventions of the `numpy` library using `numpy` arrays. You may find it helpful to start by looking at the official `numpy` Quickstart [tutorial](https://numpy.org/doc/stable/user/quickstart.html).

```python
# Run this code cell to import the NumPy library
import numpy as np
```

### Question 1
This question provides some exercise for basic indexing, slicing, and aggregating functions in `numpy`. The code below generates a 10 by 10 `numpy` array of random integers between 0 and 100 (inclusive).

```python
# Run but do not modify this code
np.random.seed(0) # so you get the same results each time
vals = np.random.randint(0, 100, (10,10)) # from 0 to 100, shape 10 by 10
print(vals)
```

Without writing any loops, write code to answer the following questions where question 1's answer is in the variable `q1_1`, question 2's answer is in `q1_2`, and so on.

1. The value at row index 2 and column index 1.
2. The last value in the first row.
3. The largest value overall.  
4. All of the values in the first column.
5. The mean of the values in the first row.
6. The minimum of the values in the last column.
7. The mean of the values in the first 5 rows and the first five columns.
8. The five smallest values in the first row, in order from least to greatest.
9. The minimum value of every row.
10. The index of the column with the largest mean value.

```python
# Write your code to answer Question 1 here
q1_1 = ...
q1_2 = ...
q1_3 = ...
q1_4 = ...
q1_5 = ...
q1_6 = ...
q1_7 = ...
q1_8 = ...
q1_9 = ...
q1_10 = ...
```

```python
grader.check("q1")
```

### Question 2
This question provides some exercise for fast element-wise computation with `numpy` ufuncs. The code below generates a random 2d `numpy` array `points` that contains 5 points, each of which has an x-value and a y-value. Each row corresponds to a point, the first column (column index 0) contains the x-value, and the second column (column index 1) contains the y-value.

```python
# Run but do not modify this code
np.random.seed(0) # so you get the same results each time
points = np.random.randn(5, 2) # draws values from standard normal distribution
print(points)
```

Without writing any loops, write code answer the following questions where quesiton 1's answer is in the variables `q2_1`, question 2's is in `q2_2`, and so on.

1. Two times every value in `points`.
2. The square of every x-value.
3. The 0-centered y-values: Subtract the mean over all y-values from every y-value.
4. The mean of the *positive* values in `points`.
5. The points with x-values greater than their y-values.
6. The magnitude of the first point. The magnitude of a point $(x, y)$ is it's Euclidean distance from the origin equal to $\sqrt{x^2 + y^2}$. 
7. The Euclidean distance between the first and the second point. The Euclidean distance formula between two points $(x_1, y_1)$ and $(x_2, y_2)$ is $\sqrt{(x_2-x_1)^2 + (y_2-y_1)^2}$.
8. The Euclidean distance between the first point and all points (you may include the distance of 0 from the first point to itself).

```python
# Write your code to answer Question 2 here
q2_1 = ...
q2_2 = ...
q2_3 = ...
q2_4 = ...
q2_5 = ...
q2_6 = ...
q2_7 = ...
q2_8 = ...
```

```python
grader.check("q2")
```

## Part 2: Pandas
In this part you will practice manipulating `pandas` Series and Dataframes. You may find it helpful to start by looking at the *10 minutes to pandas* [tutorial](https://pandas.pydata.org/pandas-docs/stable/user_guide/10min.html).

```python
# Run this code cell to import the Pandas library
import pandas as pd
```

### Question 3
The dataset `temp` imported below as a `pandas` series shows temperature deviations (in Celsius) relative to the 1951-1980 mean, courtesy of [NASA](https://climate.nasa.gov/vital-signs/global-temperature/). Note the row labels are integers (years)! Therefore, think carefully about whether an int between the square brackets is being treated as a row label or an index. 

Hint: `.loc` slices by row label while `.iloc` slices by row index.

```python
# Run but do not modify this code
temp = pd.read_csv("climate.csv", index_col=0).squeeze("columns")
print(temp.head())
temp.plot()
```

Without writing any loops, write code to answer the following questions where answers are in the variables `q3_1`, `q3_2`, and so on.

1. The temperature deviation in 1900?
2. Temperature deviations for the last three years (most recent) for which there is data
3. Temperature deviations for the last three years (most recent) for which there is data converted to Fahrenheit (temperature deviation in Fahrenheit is just 1.8 times the Celsius).
4. The mean temperature deviation since 1990 (including 1990). Remember the difference between row labels and index.
5. The number of years with positive temperature deviations.
6. The number of those years with positive temperature deviations taking place since 1970 (including 1970).
7. The five hottest years on record in the dataset.

```python
# Write your code to answer Question 3 here
q3_1 = ...
q3_2 = ...
q3_3 = ...
q3_4 = ...
q3_5 = ...
q3_6 = ...
q3_7 = ...
```

```python
grader.check("q3")
```

### Question 4
The dataset `df` imported below as a `pandas` dataframe shows statistics about nations and the happiness of citizens from the United Nations [World Happiness Report](https://worldhappiness.report).

```python
# Run but do not modify this code
df = pd.read_csv("happiness.csv", index_col = "Country name")
df.head()
```

Without writing any loops, write code to answer the following questions answers are in the variables `q4_1`, `q4_2`, and so on.

1. All of the data corresponding to the country name `United States`. 
2. All of the data corresponding to the country at row index 50.
3. The mean happiness score.
4. The number of countries with `Happiness` scores greater than 6.
5. The number of countries with `Log GDP per capita` less than 9.
6. The number of countries with `Happiness` scores greater than 6 and Log GDP per capita less than 9.

```python
# Write your code to answer Question 4 here
q4_1 = ...
q4_2 = ...
q4_3 = ...
q4_4 = ...
q4_5 = ...
q4_6 = ...
```

```python
grader.check("q4")
```

### Question 5 (Open-ended question practice)

Having now explored the data in various ways, in this question, you should analyze the data yourself by designing your own statistic. It can be an integer or float but should be a single number (as opposed to a number per state). Put the result of your analysis in the variable `statistic`. Then, explain your statistic in the written part of this question. Be precise in what you did. What data did you use? What did you not use? What kind of statistic is it? What does that value tell you about the data?

**Grading of this question**: It will be graded using an ESNU rubric as follows:
- Exemplary - Work that meets all requirements and displays full mastery of all learning goals for this question. And the code and write-up are clean and easy to understand.
- Satisfactory - Work that meets all requirements and displays at least partial mastery of all learning goals as well as full mastery of core learning goals. In other words, it is not an E, but it fulfills the requirements, including (but not necessarily only) the code matches what the write-up claims it is.
- Not Yet - Work that does not meet some requirements and/or displays developing or incomplete mastery of at least some learning goals. For example, the code and write-up do not match, or the write-up states something incorrectly.
- Unasessable - Work that is missing, does not demonstrate meaningful effort, or does not provide enough evidence to determine a level of mastery.

<!-- BEGIN QUESTION -->



```python
statistic = temp.median()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Answer 5: Written portion

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: [1, 1, 1, 1, 1, 1, 1, 1, 1, 1])
```
>>> bool(q1_1 == 37)
True
```
```
>>> bool(q1_2 == 87)
True
```
```
>>> bool(q1_3 == 99)
True
```
```
>>> q1_4 == [44, 70, 81, 47, 39, 75, 17, 35, 42, 78]
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True])
```
```
>>> import math
>>> math.isclose(q1_5, 52.5, abs_tol=0.01)
True
```
```
>>> bool(q1_6 == 12)
True
```
```
>>> import math
>>> math.isclose(q1_7, 59.16, abs_tol=0.01)
True
```
```
>>> q1_8 == [9, 21, 36, 44, 47]
array([ True,  True,  True,  True,  True])
```
```
>>> q1_9 == [9, 12, 9, 14, 9, 0, 1, 0, 3, 13]
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True])
```
```
>>> bool(q1_10 == 3)
True
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1, 1, 1])
```
>>> np.allclose(q2_1, [[3.52810469, 0.80031442], [1.95747597, 4.4817864], [3.73511598, -1.95455576], [1.90017684, -0.30271442], [-0.2064377, 0.821197]], atol=1e-05)
True
```
```
>>> np.allclose(q2_2, [3.11188068, 0.95792804, 3.48777285, 0.902668, 0.01065413], atol=1e-05)
True
```
```
>>> np.allclose(q2_3, [0.01555444, 1.85629043, -1.36188064, -0.53595997, 0.02599574], atol=1e-05)
True
```
```
>>> bool(round(q2_4, 5) == round(1.230297949608002, 5))
True
```
```
>>> np.allclose(q2_5, [[1.76405235, 0.40015721], [1.86755799, -0.97727788], [0.95008842, -0.15135721]], atol=1e-05)
True
```
```
>>> bool(round(q2_6, 5) == round(1.8088688373462207, 5))
True
```
```
>>> bool(round(q2_7, 5) == round(2.001256514017904, 5))
True
```
```
>>> np.allclose(q2_8, [0.0, 2.00125651, 1.38131852, 0.98321179, 1.86730039], atol=1e-05)
True
```


Tests for q3 (points: [1, 1, 1, 1, 1, 1, 1])
```
>>> import math
>>> math.isclose(q3_1, -0.07, abs_tol=0.01)
True
```
```
>>> def checkq32(q3_2):
...     tmp = {2017: 0.92, 2018: 0.85, 2019: 0.98}
...     flag = True
...     for i, v in q3_2.items():
...         flag = flag & (round(v, 2) == round(tmp[i], 2))
...     return flag
>>> checkq32(q3_2)
True
```
```
>>> def checkq33(q3_3):
...     tmp = {2017: 1.656, 2018: 1.53, 2019: 1.764}
...     flag = True
...     for i, v in q3_3.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq33(q3_3)
True
```
```
>>> bool(round(q3_4, 5) == round(0.5923333333333335, 5))
True
```
```
>>> bool(q3_5 == 61)
True
```
```
>>> bool(q3_6 == 46)
True
```
```
>>> def checkq37(q3_7):
...     tmp = {2018: 0.85, 2015: 0.9, 2017: 0.92, 2019: 0.98, 2016: 1.01}
...     flag = True
...     for i, v in q3_7.items():
...         flag = flag & (round(v, 2) == round(tmp[i], 2))
...     return flag
>>> checkq37(q3_7)
True
```


Tests for q4 (points: [1, 1, 1, 1, 1, 1])
```
>>> def checkq41(q4_1):
...     tmp = {'Happiness': 6.9396, 'Log GDP per capita': 10.925769, 'Social support': 0.914219, 'Healthy life expectancy': 68.2995, 'Freedom to make life choices': 0.84262, 'Generosity': 0.149892, 'Perceptions of corruption': 0.699715}
...     flag = True
...     for i, v in q4_1.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq41(q4_1)
True
```
```
>>> def checkq42(q4_2):
...     tmp = {'Happiness': 6.0218, 'Log GDP per capita': 10.340017, 'Social support': 0.93473, 'Healthy life expectancy': 68.604958, 'Freedom to make life choices': 0.877709, 'Generosity': -0.111698, 'Perceptions of corruption': 0.623074}
...     flag = True
...     for i, v in q4_2.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq42(q4_2)
True
```
```
>>> bool(round(q4_3, 5) == round(5.47323986284967, 5))
True
```
```
>>> bool(q4_4 == 53)
True
```
```
>>> bool(q4_5 == 57)
True
```
```
>>> bool(q4_6 == 5)
True
```
